# Zoekmachines

## Notebook made by

__Name__: Jesse Belleman, Levi van der Griendt, Naomi Rood, Daan van der Veen

__Student id__ : 12632295, 11399589, 12666866, 12156469

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Zoekmachines" data-toc-modified-id="Zoekmachines-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Zoekmachines</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li></ul></li><li><span><a href="#Week-1" data-toc-modified-id="Week-1-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Week 1</a></span></li><li><span><a href="#Book-Questions-(Chapter-1)" data-toc-modified-id="Book-Questions-(Chapter-1)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Book Questions (Chapter 1)</a></span><ul class="toc-item"><li><span><a href="#Exercise-1.1" data-toc-modified-id="Exercise-1.1-3.1"><span class="toc-item-num">3.1&nbsp;&nbsp;</span>Exercise 1.1</a></span></li><li><span><a href="#Exercise-1.2a" data-toc-modified-id="Exercise-1.2a-3.2"><span class="toc-item-num">3.2&nbsp;&nbsp;</span>Exercise 1.2a</a></span></li><li><span><a href="#Exercise-1.2b" data-toc-modified-id="Exercise-1.2b-3.3"><span class="toc-item-num">3.3&nbsp;&nbsp;</span>Exercise 1.2b</a></span></li><li><span><a href="#Exercise-1.3" data-toc-modified-id="Exercise-1.3-3.4"><span class="toc-item-num">3.4&nbsp;&nbsp;</span>Exercise 1.3</a></span></li><li><span><a href="#Exercise-1.4-(not-graded)" data-toc-modified-id="Exercise-1.4-(not-graded)-3.5"><span class="toc-item-num">3.5&nbsp;&nbsp;</span>Exercise 1.4 (not graded)</a></span></li><li><span><a href="#Exercise-1.5-(not-graded)" data-toc-modified-id="Exercise-1.5-(not-graded)-3.6"><span class="toc-item-num">3.6&nbsp;&nbsp;</span>Exercise 1.5 (not graded)</a></span></li><li><span><a href="#Exercise-1.6-(not-graded)" data-toc-modified-id="Exercise-1.6-(not-graded)-3.7"><span class="toc-item-num">3.7&nbsp;&nbsp;</span>Exercise 1.6 (not graded)</a></span></li><li><span><a href="#Exercise-1.7-(not-graded)" data-toc-modified-id="Exercise-1.7-(not-graded)-3.8"><span class="toc-item-num">3.8&nbsp;&nbsp;</span>Exercise 1.7 (not graded)</a></span></li><li><span><a href="#Exercise-1.8-(not-graded)" data-toc-modified-id="Exercise-1.8-(not-graded)-3.9"><span class="toc-item-num">3.9&nbsp;&nbsp;</span>Exercise 1.8 (not graded)</a></span></li><li><span><a href="#Exercise-1.9-(not-graded)" data-toc-modified-id="Exercise-1.9-(not-graded)-3.10"><span class="toc-item-num">3.10&nbsp;&nbsp;</span>Exercise 1.9 (not graded)</a></span></li></ul></li><li><span><a href="#Inverted-Index" data-toc-modified-id="Inverted-Index-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Inverted Index</a></span><ul class="toc-item"><li><span><a href="#Question-1" data-toc-modified-id="Question-1-4.1"><span class="toc-item-num">4.1&nbsp;&nbsp;</span>Question 1</a></span></li><li><span><a href="#Question-2" data-toc-modified-id="Question-2-4.2"><span class="toc-item-num">4.2&nbsp;&nbsp;</span>Question 2</a></span></li><li><span><a href="#3" data-toc-modified-id="3-4.3"><span class="toc-item-num">4.3&nbsp;&nbsp;</span>3</a></span></li><li><span><a href="#4" data-toc-modified-id="4-4.4"><span class="toc-item-num">4.4&nbsp;&nbsp;</span>4</a></span></li><li><span><a href="#5" data-toc-modified-id="5-4.5"><span class="toc-item-num">4.5&nbsp;&nbsp;</span>5</a></span></li><li><span><a href="#8-Boolean-Search" data-toc-modified-id="8-Boolean-Search-4.6"><span class="toc-item-num">4.6&nbsp;&nbsp;</span>8 Boolean Search</a></span><ul class="toc-item"><li><span><a href="#8.1-Boolean-negation-search" data-toc-modified-id="8.1-Boolean-negation-search-4.6.1"><span class="toc-item-num">4.6.1&nbsp;&nbsp;</span>8.1 Boolean negation search</a></span></li><li><span><a href="#8.2-Boolean-OR-search" data-toc-modified-id="8.2-Boolean-OR-search-4.6.2"><span class="toc-item-num">4.6.2&nbsp;&nbsp;</span>8.2 Boolean OR search</a></span></li><li><span><a href="#8.3-Ranked-boolean-search" data-toc-modified-id="8.3-Ranked-boolean-search-4.6.3"><span class="toc-item-num">4.6.3&nbsp;&nbsp;</span>8.3 Ranked boolean search</a></span><ul class="toc-item"><li><span><a href="#8.3.1-Most-words-first" data-toc-modified-id="8.3.1-Most-words-first-4.6.3.1"><span class="toc-item-num">4.6.3.1&nbsp;&nbsp;</span>8.3.1 Most-words first</a></span></li><li><span><a href="#8.3.2-Term-frequency-sum" data-toc-modified-id="8.3.2-Term-frequency-sum-4.6.3.2"><span class="toc-item-num">4.6.3.2&nbsp;&nbsp;</span>8.3.2 Term-frequency sum</a></span></li><li><span><a href="#8.3.3-Combined" data-toc-modified-id="8.3.3-Combined-4.6.3.3"><span class="toc-item-num">4.6.3.3&nbsp;&nbsp;</span>8.3.3 Combined</a></span></li></ul></li><li><span><a href="#8.4-Boolean-AND-search" data-toc-modified-id="8.4-Boolean-AND-search-4.6.4"><span class="toc-item-num">4.6.4&nbsp;&nbsp;</span>8.4 Boolean AND search</a></span></li><li><span><a href="#8.5-Ranked-Boolean-AND-search-(not-graded)" data-toc-modified-id="8.5-Ranked-Boolean-AND-search-(not-graded)-4.6.5"><span class="toc-item-num">4.6.5&nbsp;&nbsp;</span>8.5 Ranked Boolean AND search (not graded)</a></span></li><li><span><a href="#8.6-Relaxed-boolean-AND-search-(not-graded)" data-toc-modified-id="8.6-Relaxed-boolean-AND-search-(not-graded)-4.6.6"><span class="toc-item-num">4.6.6&nbsp;&nbsp;</span>8.6 Relaxed boolean AND search (not graded)</a></span></li></ul></li></ul></li><li><span><a href="#SQL-and-full-text-search" data-toc-modified-id="SQL-and-full-text-search-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>SQL and full text search</a></span></li></ul></div>

# Week 1

```python
!pip install nose
from nose.tools import assert_equal, assert_not_equal
import pandas as pd
import numpy as np
```

Output:
```
Collecting nose
  Downloading nose-1.3.7-py3-none-any.whl (154 kB)
[?25l[K     |██▏                             | 10 kB 23.7 MB/s eta 0:00:01[K     |████▎                           | 20 kB 32.4 MB/s eta 0:00:01[K     |██████▍                         | 30 kB 39.7 MB/s eta 0:00:01[K     |████████▌                       | 40 kB 28.2 MB/s eta 0:00:01[K     |██████████▋                     | 51 kB 16.2 MB/s eta 0:00:01[K     |████████████▊                   | 61 kB 17.6 MB/s eta 0:00:01[K     |██████████████▉                 | 71 kB 13.2 MB/s eta 0:00:01[K     |█████████████████               | 81 kB 14.7 MB/s eta 0:00:01[K     |███████████████████             | 92 kB 15.6 MB/s eta 0:00:01[K     |█████████████████████▏          | 102 kB 13.5 MB/s eta 0:00:01[K     |███████████████████████▎        | 112 kB 13.5 MB/s eta 0:00:01[K     |█████████████████████████▍      | 122 kB 13.5 MB/s eta 0:00:01[K     |███████████████████████████▌    | 133 kB 13.5 MB/s eta 0:00:01
```

# Book Questions (Chapter 1)
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

## Exercise 1.1

Draw the inverted index that would be built for the following document collection. (See Figure1.3 for an example.) 

**Doc 1** new home sales top forecasts 

**Doc 2** home sales rise in july 

**Doc 3** increase in home sales in july 

**Doc 4** july new home sales rise

Do it in the form of a dictionary `answer1`, with as key the word, and as value a **list** of postings

So for example: 

```
{"forecasts":[1,2,3,4],"home":[1,2,3,4],...}
```

<!-- grade_id: cell-3bac5e77cf2cbf78, points: 0, grade: false, solution: true -->
```python
answer1 = {"new":[1,4], "home":[1,2,3,4], "sales":[1,2,3,4], "top":[1], "forecasts":[1], "rise":[2,4], "in":[2,3], "july":[2,3,4], "increase":[3]}

answer1
```

Output:
```
{'forecasts': [1],
 'home': [1, 2, 3, 4],
 'in': [2, 3],
 'increase': [3],
 'july': [2, 3, 4],
 'new': [1, 4],
 'rise': [2, 4],
 'sales': [1, 2, 3, 4],
 'top': [1]}
```

<!-- grade_id: cell-07d452fd885d4634, points: 1, grade: true, solution: false -->
```python
assert_equal(type(answer1), dict)
for k,v in answer1.items():
    assert_equal(type(k), str)
    assert_equal(type(v), list)
    for p in v:
        assert_equal(type(p), int)
        assert p in range(1,5)
```

## Exercise 1.2a
Consider these documents: 

**Doc 1** breakthrough drug for schizophrenia 

**Doc 2** new schizophrenia drug 

**Doc 3** new approach for treatment of schizophrenia 

**Doc 4** new hopes for schizophrenia patients

a. Draw the term-document incidence matrix for this document collection.

Do this in pandas below. The document numbers become the index.

<!-- grade_id: cell-2f3b26c7cf356c32, points: 0, grade: false, solution: true -->
```python
data = [[1,1,1,1,1,0,0,0,0,0,0],[2,0,1,0,1,1,0,0,0,0,0],[3,0,0,1,1,1,1,1,1,0,0],[4,0,0,1,1,1,0,0,0,1,1]]


# Create the pandas DataFrame 
answer2a = pd.DataFrame(data, columns = ['document','breaktrough','drug','for','schizophrenia','new','approach','treatment','of','hopes','patients']).set_index('document')

answer2a
```

Output:
```
breaktrough  drug  for  schizophrenia  ...  treatment  of  hopes  patients
document                                         ...                                
1                   1     1    1              1  ...          0   0      0         0
2                   0     1    0              1  ...          0   0      0         0
3                   0     0    1              1  ...          1   1      0         0
4                   0     0    1              1  ...          0   0      1         1

[4 rows x 10 columns]
```

<!-- grade_id: cell-ef420142acbdf1a0, points: 0.5, grade: true, solution: false -->
```python
assert_equal(type(answer2a),pd.DataFrame)
assert_equal(set(answer2a.index), {1,2,3,4})
assert_equal(answer2a.sum().sum(),18)
```

## Exercise 1.2b
Consider these documents: 

**Doc 1** breakthrough drug for schizophrenia 

**Doc 2** new schizophrenia drug 

**Doc 3** new approach for treatment of schizophrenia 

**Doc 4** new hopes for schizophrenia patients

a. Draw the inverted index representation for this collection, as in Figure 1.3

Do it in the form of a dictionary `answer2b`, with as key the word, and as value a **list** of postings (similar to excercise 1.1).

<!-- grade_id: cell-c9191628886d80a2, points: 0, grade: false, solution: true -->
```python
answer2b = {"breaktrough":[1],"drug":[1,2],"for":[1,3,4],"schizophrenia":[1,2,3,4],"new":[2,3,4],"approach":[3],"treatment":[3],"of":[3],"hopes":[4],"patients":[4]}

answer2b
```

Output:
```
{'approach': [3],
 'breaktrough': [1],
 'drug': [1, 2],
 'for': [1, 3, 4],
 'hopes': [4],
 'new': [2, 3, 4],
 'of': [3],
 'patients': [4],
 'schizophrenia': [1, 2, 3, 4],
 'treatment': [3]}
```

<!-- grade_id: cell-539e112dcb672dcf, points: 0.5, grade: true, solution: false -->
```python
assert_equal(type(answer2b), dict)
for k,v in answer2b.items():
    assert_equal(type(k), str)
    assert_equal(type(v), list)
    for p in v:
        assert_equal(type(p), int)
```

## Exercise 1.3
For the document collection shown in Exercise 1.2, what are the returned results for these queries: 

a. schizophrenia AND drug 

b. for AND NOT(drug OR approach)

Give the answer as set of Docs (`answer3a` and `answer3b` respectively), so for example `set([1])` or `set([2,3,4])`

<!-- grade_id: cell-e1a790e2b1c4e97c, points: 0, grade: false, solution: true -->
```python
answer3a = set([1,2])
answer3b = set([4])

answer3a, answer3b
```

Output:
```
({1, 2}, {4})
```

<!-- grade_id: cell-2c885e8cc5ef265a, points: 1, grade: true, solution: false -->
```python
assert_equal(type(answer3a), set)
assert_equal(type(answer3b), set)
for a in answer3a|answer3b:
    assert_equal(type(a),int)
```

## Exercise 1.4 (not graded)
For the queries below, can we still run through the intersection in time O(x + y), where x and y are the lengths of the postings lists for Brutus and Caesar? If not, what can we achieve?

a. Brutus AND NOT Caesar 

b. Brutus OR NOT Caesar

a. Yes the time is still in O(x+y).You just check the documents that occur in the first posting list and then make sure that these documents are not in the second posting list. So you just check two posting lists.

b. Time is O(N), with N being the total amount of documents. This is because it tells us "OR NOT" ceasar. This means that there can possibly be no ceasar in all documents so in that case you would need to write down all document numbers instead of just the posting list behind ceasar. 

## Exercise 1.5 (not graded)
Extend the postings merge algorithm to arbitrary Boolean query formulas. What is its time complexity? For instance, consider:

c. (Brutus OR Caesar) AND NOT (Antony OR Cleopatra)

Can we always merge in linear time? Linear in what? Can we do better than this?


ANSWER: 
We merge in O(a N). a are the amount of terms we use in the query and N is the amount of documents. Therefore the intersection time is linear in terms of N and a. So in terms of the amount of document and the amount of terms used in the query. You can not do better than this because N is the amount of documents and these can all be in the result (for example if Antony OR Cleopatra are in no documents, the right side of the query becomes as big as N).

## Exercise 1.6 (not graded)
We can use distributive laws for AND and OR to rewrite queries.

a. Show how to rewrite the query in Exercise1.5 into disjunctive normal form using the distributive laws. 

b. Would the resulting query be more or less efﬁciently evaluated than the original formof this query? 

c. Is this result true in general or does it depend on the words and the contents of the document collection?


ANSWER:
a.
(Brutus OR Caesar) AND NOT (Antony OR Cleopatra) = 


(Brutus AND NOT (Antony OR Cleopatra)  OR (Caesar AND NOT (Antony OR Cleopatra) = 

(Brutus AND NOT Antony AND NOT Cleopatra ) OR (Caesar AND NOT Antony AND NOT Cleopatra)


b. This query is more efficient that the query in exc 1,5. This is because in exc 1,5 we start with two unions. These unions can become pretty big lists and then checking both these lists for intersections can take a long time. Whereas the newly written query starts with intersections on the smaller posting lists (instead of on the bigger created unions, like in the query of exc 1,5) to then only form a union of these lists. 

## Exercise 1.7 (not graded)
Recommend a query processing order for<br/>
d. (tangerine OR trees) AND (marmalade OR skies) AND (kaleidoscope OR eyes) given the following postings list sizes:<br/>

|Term |Postings size|
|-----|:------------|
|eyes|213312|
|kaleidoscope|87009|
|marmalade|107913|
|skies|271658|
|tangerine|46653|
|trees|316812|


ANSWER: 1 = tangerine OR trees, which costs O(46653+316812) = O(363465)

2 = marmalade OR skies, which costs O(107913+271658) = O(379571)

3 = kaleidoscoop OR eyes, which costs O(87009 + 213312 = O(300321)

It is therefore better to first merge the lists: (tangerine OR trees) AND (kaleidoscoop OR eyes) and then merge this result with the list (marmalade OR skies). This is because the shorter the lists during the AND operation the less time it takes and the most likely that the result of the intersection will be smaller. A smaller result will result in less time needed to do an intersection with the resulting list.

## Exercise 1.8 (not graded)
If the query is:<br/>

e. friends AND romans AND (NOT countrymen)<br/>

how could we use the frequency of countrymen in evaluating the best query evaluation order? In particular, propose a way of handling negation in determining the order of query processing.

ANSWER: We know the amount of documents N. And we can check the size of the posting list of countrymen (s) and then take N-s to determine the query evaluation order.

## Exercise 1.9 (not graded)
For a conjunctive query, is processing postings lists in order of size guaranteed to be optimal? Explain why it is, or give an example where it isn’t.

ANSWER: It is not always guaranteed to be optimal. This is because you can not know the size of the intersections on forehand. Some big lists of 1000 words might have only a dozen words in common whilst one of these 1000 words list might have 100 words in common with a smaller list of 500 words. 

# Inverted Index

<!-- grade_id: cell-480ae65e69b5d934, points: 0, grade: true, solution: false -->
```python
%matplotlib inline
from collections import Counter, defaultdict, OrderedDict
from bs4 import BeautifulSoup
import os
from tqdm import tqdm_notebook
import glob
import nltk
import zipfile
import math
import pandas as pd
import sys
import itertools


def loadShakespeare():
    if 'shaks200.zip' in os.listdir():
        return 'shaks200.zip'
    elif os.path.exists('../../data/Week1/'):
        return '../../data/Week1/shaks200.zip'
    elif os.path.exists('../../../data/Week1/'):
        return '../../../data/Week1/shaks200.zip'

    
```

## Question 1
Make a function `index_collection` which turns a zipfolder (like Shakespeare) into an inverted index.

So the index will be of form:
```
{'the': defaultdict(int,
                         {'a_and_c': 875,
                          'all_well': 736,
                          'as_you': 698,
                          'com_err': 445,
                          'coriolan': 1130,
                          'cymbelin': 973,
                          'dream': 565,
                          'hamlet': 1146,
                          ...
```

* Think about lower case, tokenization etc.
* Doesn't have to be a defaultdict, but can be useful.

<!-- grade_id: cell-ecc0e099ac33f896, points: 1, grade: false, solution: true -->
```python
from tqdm import notebook
import re
from nltk.tokenize import word_tokenize

nltk.download("punkt")

def index_collection(zipfolder):
    # With zipfile we can read the file without opening the zip file
    archive = zipfile.ZipFile(zipfolder, 'r')
    namelist = [x for x in archive.namelist() if '.xml' in x]
    MyIndex = defaultdict(lambda: defaultdict(int)) # initialize MyIndex
    for infile in notebook.tqdm(namelist): # loop over each file
        f = archive.open(infile)
        file_name = infile.replace('.xml', '')
        for line in f.readlines():
            # decode line from byte to string object
            line = line.decode("utf-8")
            # search only for tet between two xml tags
            line = re.search('<.+>(.+?)</.+>', line)
            # if there was any text between the tags
            if line:
                # select between the tags and lowercase it
                line = line.group(1)
                line = line.lower()
                # forms a list of the spoken words
                line = word_tokenize(line)
                
                # iterate over the words in the line
                for word in line:
                    # check if word is a word not a punctuation mark
                    if word[0].isalnum():

                      # check if the word has already been mentioned in this play
                      if file_name in MyIndex[word]:
                          # then increment the count for this word in this play with 1
                          MyIndex[word][file_name] += 1   
                      # else add the mentioning of this word in this play to the dict
                      else:
                          MyIndex[word][file_name] = 1 
                        
    return MyIndex

%time Shakespeare = index_collection(loadShakespeare())


Shakespeare['the'], Shakespeare['witch']
```

Output:
```
[nltk_data] Downloading package punkt to /root/nltk_data...
[nltk_data]   Package punkt is already up-to-date!
/usr/lib/python3.7/zipfile.py in _RealGetContents(self)
   1319         fp = self.fp
   1320         try:
-> 1321             endrec = _EndRecData(fp)
   1322         except OSError:
   1323             raise BadZipFile("File is not a zip file")

/usr/lib/python3.7/zipfile.py in _EndRecData(fpin)
    257 
    258     # Determine file size
--> 259     fpin.seek(0, 2)
    260     filesize = fpin.tell()
    261 

AttributeError: 'NoneType' object has no attribute 'seek'
```

<!-- grade_id: cell-80a890412656a8d0, points: 1, grade: true, solution: false -->
```python
assert isinstance(Shakespeare, dict)
assert 'macbeth' in Shakespeare['hurlyburly']

```

## Question 2

Count the total number of tokens in the works of Shakespeare (`TotNumberofTokens`) and the total number of unique tokens (`Vocsize`).  Answer with two one-liners using your created index (which is called `Shakespeare`).

<!-- grade_id: cell-0d0ac196213493f2, points: 0, grade: false, solution: true -->
```python
# hint: look up how you compute the sum of a list of numbers without a for loop or an extra variable
Vocsize = len(Shakespeare.keys())
TotNumberofTokens = sum([tot for words in  Shakespeare.values() for tot in words.values()])

Vocsize, TotNumberofTokens, round(TotNumberofTokens/Vocsize,1)
```

<!-- grade_id: cell-280ef3963b3a3ff0, points: 0.5, grade: true, solution: false -->
```python
assert_equal(type(Vocsize), int)
```

<!-- grade_id: cell-543ceebff67e59c4, points: 0.5, grade: true, solution: false -->
```python
assert_equal(type(TotNumberofTokens), int)
```

## 3

1. Calculate the document frequency (in how many documents occurs the term) for each term (store this in the dictionary named `DocFreq`) and 
2. the corpus frequency (how often does the term occur in the corpus) for each term (store this in the dictionary named `CorpusFreq`). 
3. Add these two values as new values in your inverted index, which will be called `NewIndex`.

The `NewIndex` will then look like this:
```{'the': {'CorpusFreq': 28254,
  'DocFreq': 37,
  'Freq_per_Doc': Counter({'a_and_c': 874,
           'all_well': 736,
           'as_you': 698,
           'com_err': 444,
           'coriolan': 1130,
           'cymbelin': 973,
           'dream': 565,
           'hamlet': 1146,
           'hen_iv_1': 869,
           'hen_iv_2': 1001,
           'hen_v': 1089,
           'hen_vi_1': 731,
           'hen_vi_2': 953,
           'hen_vi_3': 818,
           'hen_viii': 947,
           'j_caesar': 607,
           'john': 734,
           'lear': 911,
           'lll': 860,
           'm_for_m': 702,
           'm_wives': 610,
           'macbeth': 736,
           'merchant': 840,
           'much_ado': 581,
           'othello': 765,
           'pericles': 635,
           'r_and_j': 687,
           'rich_ii': 752,
           'rich_iii': 995,
           't_night': 550,
           'taming': 510,
           'tempest': 521,
           'timon': 510,
           'titus': 659,
           'troilus': 846,
           'two_gent': 409,
           'win_tale': 860})},...```
           
 * the type of `Freq_per_Doc` doesn't have to be a Counter, can also be a defaultdict or dict.

<!-- grade_id: cell-739b18d157f01e44, points: 0, grade: false, solution: true -->
```python

CorpusFreq = {word:sum(Shakespeare[word].values()) for word in Shakespeare.keys()} # change to your solution
DocFreq = {word:len(Shakespeare[word].keys()) for word in Shakespeare.keys()} # change to your solution
NewIndex = { word: {'CorpusFreq': CorpusFreq[word], "DocFreq": DocFreq[word],
            "Freq_per_Doc": {doc: Shakespeare[word][doc] for doc in Shakespeare[word].keys()}} 
            for word in Shakespeare.keys() } # change to your solution
 


NewIndex['the']
```

<!-- grade_id: cell-4cdc200e7c434c2b, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(CorpusFreq),dict)
```

<!-- grade_id: cell-9807e00ccea6371d, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(DocFreq),dict)
```

<!-- grade_id: cell-dcf52b08a1da000e, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(NewIndex),dict)
assert_equal(set(NewIndex['the'].keys()),{'Freq_per_Doc','DocFreq','CorpusFreq'})

```

## 4

Create a function `IndexInfo(index)` which from the enhanced index retrieves the following information:

* size of the vocabulary
* nr of terms with corpus frequency 1
* nr of terms with document frequency 1 
* nr of terms with document frequency  equal to the number of documents in the corpus
* nr of terms with document frequency  half or more than the number of documents in the corpus

It should return the information in that order!

<!-- grade_id: cell-30fe60f6fba5acab, points: 0, grade: false, solution: true -->
```python
def documents_to_list(doclist,index):
    for doc in index.keys():
      if doc not in doclist:
        doclist.append(doc);
    return doclist

def IndexInfo(index):
    documents = []
    half = {}
    full = {}
    Vocsize = 0
    Hapaxes = 0
    SingleDocTerms = 0
    FullDocTerms = 0
    HalfDocTerms = 0

    # we going walk though the words
    for word in index.keys():
      Vocsize = Vocsize + 1
      # check if it has a corpus frequency of 1 also have a document frequency of 1
      if index[word]['CorpusFreq'] == 1:
        Hapaxes = Hapaxes + 1
        SingleDocTerms = SingleDocTerms + 1
      # check if the is only 1 document.
      elif index[word]['DocFreq'] == 1:
        SingleDocTerms = SingleDocTerms + 1
      # add all document this documents in is.
      documents = documents_to_list(documents,index[word]['Freq_per_Doc'])
      # add possible 
      if index[word]["DocFreq"] == len(documents):
        full[word] = index[word]["DocFreq"]
        half[word] = index[word]["DocFreq"]
      elif index[word]["DocFreq"] == len(documents)/2:
        half[word] = index[word]["DocFreq"]
    
    #now check if they all possible half and full length are acctual full and half.
    for word  in full.keys():
      if len(documents) == full[word]:
        FullDocTerms = FullDocTerms + 1
    #here we do the same as above only for half or more
    halfsize = len(documents)/2
    for word in half.keys():
      if halfsize <= half[word]:
        HalfDocTerms = HalfDocTerms + 1
      
      
    

    return Vocsize, Hapaxes, SingleDocTerms, FullDocTerms, HalfDocTerms
IndexInfo(NewIndex)
```

<!-- grade_id: cell-ca98643b040c6eac, points: 0.2, grade: true, solution: false -->
```python
assert_equal(type(IndexInfo(NewIndex)[0]), int)
```

<!-- grade_id: cell-322261b39549eb9b, points: 0.2, grade: true, solution: false -->
```python
assert_equal(type(IndexInfo(NewIndex)[1]), int)
```

<!-- grade_id: cell-301745b87a6b9806, points: 0.2, grade: true, solution: false -->
```python
assert_equal(type(IndexInfo(NewIndex)[2]), int)
```

<!-- grade_id: cell-03e46e9370e4d44d, points: 0.2, grade: true, solution: false -->
```python
assert_equal(type(IndexInfo(NewIndex)[3]), int)
```

<!-- grade_id: cell-f5e4c181b23654c6, points: 0.2, grade: true, solution: false -->
```python
assert_equal(type(IndexInfo(NewIndex)[4]), int)
```

## 5

Look up what Zipf's law means in Wikipedia or in the book. Now show that it holds or not for our corpus.

Make the pandas dataframe  `CorpusFreqDF` so that `CorpusFreqDF.plot(loglog=True);` will return a Zipf plot of terms and frequencies.

<!-- grade_id: cell-b201e9f17c20805f, points: 0, grade: false, solution: true -->
```python
CorpusFreqDF = pd.DataFrame([NewIndex[word]['CorpusFreq'] for word in NewIndex], columns=['CorpusFrequency'])
CorpusFreqDF = CorpusFreqDF.sort_values(by=['CorpusFrequency'], ascending=False)
CorpusFreqDF = CorpusFreqDF['CorpusFrequency'].reset_index()
CorpusFreqDF = CorpusFreqDF['CorpusFrequency'].to_frame()

CorpusFreqDF.plot(loglog=True);
```

<!-- grade_id: cell-bb71005d7439ea25, points: 1, grade: true, solution: false -->
```python
assert_equal(type(CorpusFreqDF), pd.DataFrame)
CorpusFreqDF.shape,CorpusFreqDF.CorpusFrequency.values
```

## 8 Boolean Search

We will now make a Boolean search engine as described in Chapter 1.

### 8.1 Boolean negation search
Create the function `booleanSearchNOT(word, index)`

<!-- grade_id: cell-ee3349180eacd8da, points: 0, grade: false, solution: true -->
```python
def booleanSearchNOT(word, index):
    """
    This function takes a string as input and returns all document ids as set
    which do NOT contain the word.
    """
   
    # set of all documents
    doc_list = set()
    for key in index.keys():
        for doc in index[key]['Freq_per_Doc']:
            doc_list.add(doc)    
    
    # set of all the documents that the word DOES occur in.
    occurs = set(index[word]['Freq_per_Doc'].keys())
    
    return doc_list - occurs
    
booleanSearchNOT("lightly", NewIndex)
```

<!-- grade_id: cell-f6125c2504b1a68f, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchNOT("lightly", NewIndex)), set)
```

### 8.2 Boolean OR search

Create the function `booleanSearchOR(query, index)`

<!-- grade_id: cell-10aa5d1597de0be5, points: 0, grade: false, solution: true -->
```python
def booleanSearchOR(query, index):
    """
    This function takes a query as input and returns all document ids as set which
    contain at least one of the words.
    """
    result = set()
    for word in query:
        for document in index[word]['Freq_per_Doc']:
            result.add(document)
            
    return result
    
    
booleanSearchOR(["lightly", "love"], NewIndex)
```

<!-- grade_id: cell-f8efc9467f05480f, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchOR(["lightly", "love"], NewIndex)), set)
```

### 8.3 Ranked boolean search
A Boolean search engine is not supposed to rank the documents it returns, but we can do that anyway.

#### 8.3.1 Most-words first
Rank the documents returned for the OR-query in last question as follows: documents which match most words are returned first (if equal, sort alphabetically on document ID). Return the document ID's followed by the number of matching words (i.e. the "score" of the document for this query).

<!-- grade_id: cell-facf62845225258f, points: 0, grade: false, solution: true -->
```python
def booleanSearchMatch(query, index):
    """
    This function takes a query as input and returns a list of  all document ids and
    scores of documents which contain at least one of the words. Documents
    are ranked such that documents that match the most words are returned
    first.
    """
    
    ranked = {}
    
    # add the documents with their count to a dict
    for word in query:
        for document in index[word]['Freq_per_Doc']:
            if document in ranked.keys():
                ranked[document] += 1
            else:
                ranked[document] = 1
    
    # sort the document on the count
    ranked = sorted(ranked.items(), key=lambda x: x[1], reverse=True)

    return ranked
    
booleanSearchMatch(["love","lightly"], NewIndex)
```

<!-- grade_id: cell-0f1a4c699d049875, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchMatch(["love","lightly"], NewIndex)), list)
assert_equal(type(booleanSearchMatch(["love","lightly"], NewIndex)[0]), tuple)
assert_equal(type(booleanSearchMatch(["love","lightly"], NewIndex)[0][0]), str)
assert_equal(type(booleanSearchMatch(["love","lightly"], NewIndex)[0][1]), int)
```

#### 8.3.2 Term-frequency sum
Now rank the documents by the sum of the term frequencies for all query words. Return doc ID's and scores.

Again first sort by score, and break ties by sorting  alphabetically on Doc ID.

<!-- grade_id: cell-74c7e3f7e51b6c92, points: 0, grade: false, solution: true -->
```python
def booleanSearchSumTF(query, index):
    """
    This function takes a query as input and returns all document ids and
    scores of documents which contain at least one of the words. Documents
    are ranked by the sum of the term frequencies of all query words.
    """
    
    ranked = {}
    
    # add the documents with their count to a dict
    for word in query:
        for document in index[word]['Freq_per_Doc']:
            if document in ranked.keys():
                ranked[document] += index[word]['DocFreq']
            else:
                ranked[document] = index[word]['DocFreq']
    
    # sort the document on the count
    ranked = sorted(ranked.items(), key=lambda x: x[1], reverse=True)

    return ranked
    
booleanSearchSumTF(["love","lightly"], NewIndex)
```

<!-- grade_id: cell-83ff086299885d29, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchSumTF(["love","lightly"], NewIndex)), list)
assert_equal(type(booleanSearchSumTF(["love","lightly"], NewIndex)[0]), tuple)
assert_equal(type(booleanSearchSumTF(["love","lightly"], NewIndex)[0][0]), str)
assert_equal(type(booleanSearchSumTF(["love","lightly"], NewIndex)[0][1]), int)
```

#### 8.3.3 Combined
Now combine the two rankings. First rank on number of matching search terms, and then on the sum of the term counts. Return both scores for each DocID.

The tuple will thus look like this: `(document ID, number of matching search terms, sum of the term frequency counts)`

<!-- grade_id: cell-13a89b3b12fd9a1d, points: 0, grade: false, solution: true -->
```python
def booleanSearchCombined(query, index):
    """
    This function takes a query as input and returns all document ids and
    scores of documents which contain at least one of the words. Documents
    are ranked first on the number of matching search terms, and then on
    the sum of the term frequency counts.
    """
    
    ranked = []
    
    # used to keep track of the location of a document in the ranked list
    in_list = {}
    counter = 0
    
    # add the documents with their Doc frequency's to a dict
    for word in query:
        for document in index[word]['Freq_per_Doc']:
            # if document was already in the ranked list
            if document in in_list:
                # increment its values
                ranked[in_list[document]][1] += 1
                ranked[in_list[document]][2] += index[word]['DocFreq']
            # else we add it to the ranked list
            else:
                in_list[document] = counter
                counter += 1
            
                ranked.append([document, 1, index[word]['DocFreq']])
                
 
    # sort the document (the minus one before x[1] and x[2] are needed to reverse the sorting of the count
    # but to keep the alphabetic sorting order)
    ranked = sorted(ranked, key=lambda x: (-x[1],-x[2], x[0]))
    
    # convert to tuple
    for i in range(len(ranked)):
        ranked[i] = tuple(ranked[i])

    return ranked   
    
booleanSearchCombined(["love","lightly"], NewIndex)
```

<!-- grade_id: cell-6cdd92cfcc2d415d, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchCombined(["love","lightly"], NewIndex)), list)
assert_equal(type(booleanSearchCombined(["love","lightly"], NewIndex)[0]), tuple)
assert_equal(type(booleanSearchCombined(["love","lightly"], NewIndex)[0][0]), str)
assert_equal(type(booleanSearchCombined(["love","lightly"], NewIndex)[0][1]), int)
assert_equal(type(booleanSearchCombined(["love","lightly"], NewIndex)[0][2]), int)
```

### 8.4 Boolean AND search
Create the function `booleanSearchAND(query, index)`

<!-- grade_id: cell-e507e9a68502f0e8, points: 0, grade: false, solution: true -->
```python
def booleanSearchAND(query, index):
    """
    This function takes a query as input and returns all document ids as set which
    contain ALL the words.
    """
    
    sets = []
    
    # iterate over all the words in the query
    for word in query:
        # create a set for this particular word
        word_set = set()
        # add al the documents in which this word occurs to its set
        for document in index[word]['Freq_per_Doc']:
            word_set.add(document)
        # add the set of documents for this particular word to the total list of sets
        sets.append(word_set)
  
    # return the intersection of all the sets in the list of sets
    return sets[0].intersection(*sets)
    
booleanSearchAND(["lightly", "love"], NewIndex)
```

<!-- grade_id: cell-ab54ede8ff89e3c7, points: 1, grade: true, solution: false -->
```python
assert_equal(type(booleanSearchAND(["lightly", "love"], NewIndex)), set)
```

### 8.5 Ranked Boolean AND search (not graded)
How would you implement a ranking for the Boolean AND?

You can implement it for the biggest part like the normal Boolean AND search. But to rank all the documents, it is possible to sum the number of every time that a word of the query is in the document. So the different words in a query will count for the ranking.

### 8.6 Relaxed boolean AND search (not graded)
How can you implement a "relaxed version" of Boolean AND, in which you also give results if not ALL query words are present, but as much as possible? Of course you want to rank the documents again.

To implement relaxed boolean AND search it is possible to do this to store different information about the documents. Make a dict with the documents as keys and the values represent a list with 2 items. The first item is the amount of words of the query are in the document and the second item is the sum of all words that occures in the document. Loop trough the documents and add a document to the dict if a word of the query exists. Last, rank the documents first by the number of words of the query that are in the document and second by the sum of all words that occures in the document.


# SQL and full text search
 
 Many SQL database systems allow full text search on columns storing text. Unlike with SQL the functionality is not that much standardized but it follows similar ideas. 
 
 In this exercise you create a bollean search engine, possibly with ranking, in which you can combine full text search with restrictions on other fields. The database used is a collection of transcripts of TedX talks.
 
See the notebook `SQLLiteFullText.ipynb` in this folder for code to get you started and the full description of th exercise. 

Simply paste your answer to this question here. 