# Zoekmachines

## Notebook made by

__Name__|Zirk| en vrienden
__Student id__ : 

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Zoekmachines" data-toc-modified-id="Zoekmachines-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Zoekmachines</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li></ul></li><li><span><a href="#Week-5" data-toc-modified-id="Week-5-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Week 5</a></span></li><li><span><a href="#Evaluation-of-Document-Classification-(4pt)" data-toc-modified-id="Evaluation-of-Document-Classification-(4pt)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Evaluation of Document Classification (4pt)</a></span></li><li><span><a href="#Relevance-feedback" data-toc-modified-id="Relevance-feedback-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Relevance feedback</a></span></li><li><span><a href="#Book-excercises-(Chapter-12)" data-toc-modified-id="Book-excercises-(Chapter-12)-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Book excercises (Chapter 12)</a></span><ul class="toc-item"><li><span><a href="#Exercise-12.3" data-toc-modified-id="Exercise-12.3-5.1"><span class="toc-item-num">5.1&nbsp;&nbsp;</span>Exercise 12.3</a></span></li><li><span><a href="#Exercise-12.4" data-toc-modified-id="Exercise-12.4-5.2"><span class="toc-item-num">5.2&nbsp;&nbsp;</span>Exercise 12.4</a></span></li><li><span><a href="#Exercise-12.5" data-toc-modified-id="Exercise-12.5-5.3"><span class="toc-item-num">5.3&nbsp;&nbsp;</span>Exercise 12.5</a></span></li><li><span><a href="#Exercise-12.6" data-toc-modified-id="Exercise-12.6-5.4"><span class="toc-item-num">5.4&nbsp;&nbsp;</span>Exercise 12.6</a></span></li><li><span><a href="#Exercise-12.7" data-toc-modified-id="Exercise-12.7-5.5"><span class="toc-item-num">5.5&nbsp;&nbsp;</span>Exercise 12.7</a></span></li><li><span><a href="#Exercise-12.8" data-toc-modified-id="Exercise-12.8-5.6"><span class="toc-item-num">5.6&nbsp;&nbsp;</span>Exercise 12.8</a></span></li><li><span><a href="#Exercise-12.9" data-toc-modified-id="Exercise-12.9-5.7"><span class="toc-item-num">5.7&nbsp;&nbsp;</span>Exercise 12.9</a></span></li></ul></li></ul></div>

# Week 5

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

# Evaluation of Document Classification (4pt)


Stel je hebt een Naive Bayes Document classifier getrained voor 1 klasse $C$. Stel je hebt een test-collectie van 400 documenten. 

Beantwoord de volgende vragen in de cellen hieronder.

1. De classifier heeft elk document in de collectie geklassificeerd. Wat betekent het als je een precisie hebt van 60%? Hoeveel documenten zijn dan onterecht als behorende tot klasse $C$ geklassificeerd?
2. Stel dat 300 van de test documenten in klasse $C$ vallen? Welke precisie baseline hoort bij deze test-collectie? Welke recall baseline? 
3. Stel je classifier heeft een recall van 50%. Stel 100 van de 400 docs zitten in klasse $C$.  De classifier hoeft niet elk document te classificeren.   Geef exact   aan wat de precisie kan zijn voor deze classifier. Geef dus de onder en bovengrens voor de precisie. 
4. Stel 100 van de 400 docs zitten in klasse $C$. Welke baseline heb je voor _accurracy_?

Met een classificatie baseline bedoelen we een classifier die helemaal niet naar de input kijkt maar alleen slim kiest op basis van de verdeling van de documenten over de klasse $C$.
Denk bijvoorbeeld aan de vraag naar het geslacht van een willekeurige KI student. Wat zou je kiezen als je niks van die student weet? Wat is de recall en de precisie voor het geslacht "jongen" dan?

<!-- grade_id: dc1-a, points: 0, grade: false, solution: true -->
```python
# deel vraag 1
onterecht = (1 - 0.6) * 400  # = 40% van 400 = 160
onterecht

```

Output:
```
160.0
```

<!-- grade_id: dc1-b, points: 0, grade: false, solution: true -->
```python
# deel vraag 2
# classificeer alles als klasse C
precisie_baseline = 300 / 400  # 300 / 400 =  75% = 0.75
recall_baseline = 300 / 300    # 300 / 300 = 100% = 1

```

<!-- grade_id: dc1-c, points: 0, grade: false, solution: true -->
```python
# deel vraag 3
ondergrens = (0.5 * 100) / (400 - 100 + 0.5 * 100)  # alle andere verkeerde documenten
bovengrens = 1   # geen verkeerde documenten

```

<!-- grade_id: dc1-d, points: 0, grade: false, solution: true -->
```python
# deel vraag 4
# classificeer alles als niet klasse C
accuracy_baseline = 300 / 400

```

<!-- grade_id: dc1at, points: 1, grade: true, solution: false -->
```python
assert 0<=onterecht<= 400


```

<!-- grade_id: dc1bt, points: 1, grade: true, solution: false -->
```python
assert 0<=precisie_baseline<=1
assert 0<=recall_baseline<= 1

```

<!-- grade_id: dc1ct, points: 1, grade: true, solution: false -->
```python
assert 0<=ondergrens<=1
assert 0<=bovengrens<= 1
```

<!-- grade_id: dc1dt, points: 1, grade: true, solution: false -->
```python
assert 0<=accuracy_baseline<=1

```

# Relevance feedback

In deze vraag gaan we de SMART versie van Rochio relevance feedback (vergelijking 9.3 in het boek) implementeren in Python. We gaan uit van een collectie met een vocabulair van 100 termen. Elke query en elk document is gerepresenteerd als een vector van lengte 100 gevuld met getallen tussen 0 en 1. 

Maak de functies voor `centroid` en `modified_query` af. Alle data zit in numpy arrays, dus je kunt er heel makkelijk mee werken op een manier die je kent uit linaire algebra. 

De data bestaat uit 
* een query van 2 woorden
* 8 relevante documenten en 3 niet relevante documenten

Hint: beide functies kan je zonder for loops schrijven, zelfs comprehensies zijn niet nodig.

```python
# data
q=np.zeros(100)
q[8],q[47]= .4, .111
Relevant_docs=np.array([np.random.rand(100) for _ in range(8)])
NonRelevant_docs=np.array([np.random.rand(100) for _ in range(3)])

print(q)
```

Output:
```
[0.    0.    0.    0.    0.    0.    0.    0.    0.4   0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.111
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.   ]
```

<!-- grade_id: rfa, points: 0, grade: false, solution: true -->
```python
def centroid(D):
    return D.mean(axis=0)

# test
print(centroid(NonRelevant_docs ))

def modified_query(query,relevantdocs,nonrelevantdocs,alpha=1,beta=.75,gamma=.25):
    return alpha * query + beta * centroid(relevantdocs) - gamma * centroid(nonrelevantdocs)

# test
modified_query(q,Relevant_docs,NonRelevant_docs)

```

Output:
```
[0.89844379 0.68214989 0.51960338 0.67325297 0.66315413 0.66297871
 0.69481716 0.33761536 0.6675948  0.71770857 0.50292111 0.56871138
 0.75189894 0.34270258 0.56629882 0.43764204 0.18711292 0.32299684
 0.27838118 0.59070851 0.26367413 0.81308736 0.43036037 0.28271928
 0.69526472 0.18674273 0.84077527 0.5856009  0.31614511 0.63306211
 0.50078868 0.72199574 0.43570542 0.60768832 0.51932455 0.88865269
 0.78736268 0.53521806 0.74529133 0.41219506 0.60175989 0.43146614
 0.64090355 0.6435868  0.25018276 0.42928504 0.71298853 0.59961288
 0.49004699 0.80022838 0.6567588  0.50346696 0.12642083 0.61176966
 0.65668375 0.55978588 0.616008   0.67067375 0.82980108 0.35798093
 0.4674791  0.81685053 0.36546271 0.39408138 0.32185955 0.38818848
 0.58216815 0.31394017 0.24691462 0.53505759 0.6799137  0.5519244
 0.41688279 0.43970197 0.25140122 0.54487723 0.54167915 0.42577319
 0.82159769 0.43827168 0.52583991 0.54971897 0.32853396 0.30610317
 0.50316729 0.45774636 0.41532061 0.48867454 0.46884626 0.31191
```

<!-- grade_id: rft, points: 1, grade: true, solution: false -->
```python
assert isinstance(modified_query(q,Relevant_docs,NonRelevant_docs), np.ndarray)

```

# Book excercises (Chapter 12)
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

## Exercise 12.3
What  is  the  likelihood  ratio  of  the  document  according  to M1 and M2 in  Example 12.2?

<!-- grade_id: cell-1078c13a222d6f56, points: 0, grade: false, solution: true -->
```python
likelihood_ratio = 0.00000000000048 / 0.000000000000000384

likelihood_ratio
```

Output:
```
1250.0
```

<!-- grade_id: cell-d5b4ca4da32d2e9c, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_ratio) in [int, float]
```

## Exercise 12.4
No explicit STOP probability appeared in Example 12.2. Assuming that the STOP probability of each model is 0.1, does this change the likelihood ratio of a document according to the two models?

<!-- grade_id: cell-88d84f5f3e274f8c, points: 0, grade: false, solution: true -->
```python
# Doesn't change, because if the chance of STOP is equal, both are multiplied by 0.1
# Multiplying the numerator and denominator of a fraction by the same amount doesn't change the result
likelihood_with_stop = (0.00000000000048 * 0.1) / (0.000000000000000384 * 0.1)

likelihood_with_stop
```

Output:
```
1250.0
```

<!-- grade_id: cell-77b507102f5e246e, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_with_stop) in [int, float]
```

## Exercise 12.5
How might a language model be used in a spelling correction system? In particular, consider the case of context-sensitive spelling correction, and correcting incorrect usages of words, such as _their_ in _Are you their?_(See Section3.5(page65) for pointers to some literature on this topic.)

<!-- grade_id: cell-4bf9772cefb54700, points: 1, grade: true, solution: true -->
```python
"""
If the chance of “Are you there” is significantly higher than the
chance of “Are you their” than we might suggest a correction of
“their” into “there”.
"""
```

Output:
```
'\nIf the chance of “Are you there” is significantly higher than the\nchance of “Are you their” than we might suggest a correction of\n“their” into “there”.\n'
```

## Exercise 12.6
Consider making a language model from the following training text:
    the martian has landed on the latin pop sensation ricky martin
    
a. Under a MLE-estimated unigram probability model, what are P(the) and P(martian)?

b. Under a MLE-estimated bigram model, what are P(sensation|pop) and P(pop|the)?

<!-- grade_id: cell-ecc32bc65a9fbd20, points: 1, grade: false, solution: true -->
```python
from collections import Counter

text = "the martian has landed on the latin pop sensation ricky martin".split()
counts = Counter(text)
total = len(text)      # = sum(counts.values())

P_the = counts['the'] / total          # text.count("the") / total
P_martian = counts['martian'] / total  # text.count("martian") / total
P_sensation_pop = 1  # all (one) occurences of pop followed by sensation
P_pop_the = 0        # no occurence of pop followed by the

P_the, P_martian, P_sensation_pop, P_pop_the
# len(text),text.count("the"), text.count("martian")
```

Output:
```
(0.18181818181818182, 0.09090909090909091, 1, 0)
```

<!-- grade_id: cell-4bf550a6e30c9cce, points: 0.25, grade: true, solution: false -->
```python
assert type(P_the) in [int, float]
```

<!-- grade_id: cell-cf1055abc4db1355, points: 0.25, grade: true, solution: false -->
```python
assert type(P_martian) in [int, float]
```

<!-- grade_id: cell-c7730c5c9cd22dc6, points: 0.25, grade: true, solution: false -->
```python
assert type(P_sensation_pop) in [int, float]
```

<!-- grade_id: cell-e1b617addf2268f7, points: 0.25, grade: true, solution: false -->
```python
assert type(P_pop_the) in [int, float]
```

## Exercise 12.7
Suppose we have  a collection that consists of the 4 documents given in the below table.

| docID | Document text                              |
|-------|--------------------------------------------|
| 1     | click go the shears boys click click click |
| 2     | click click                                |
| 3     | metal here                                 |
| 4     | metal shears click here                    |


Build a  query likelihood language  model for  this  document  collection. Assume  amixture model between the documents and the collection, with both weighted at 0.5. Maximum likelihood estimation (mle) is used to estimate both as unigram models. Work out the model probabilities of the queries `click`, `shears`, and hence `click shears` for each document, and use those probabilities to rank the documents returned by eachquery.

For each query give the probablities in a list, where the first value corresponds to Doc1, the second value to Doc2 etc.

<!-- grade_id: cell-5a45af37333b7b02, points: 0, grade: false, solution: true -->
```python
click_chance = [4/8, 2/2, 0/2, 1/4]
collection_click_chance = (4+2+0+1) / (8+2+2+4)
shears_chance = [1/8, 0/2, 0/2, 1/4]
collection_shears_chance = (1+0+0+1) / (8+2+2+4)

# Assuming a uniform prior, we can ignore P(d)...
P_click = [0.5 * click_chance[i] + 0.5 * collection_click_chance for i in range(4)]
P_shears = [0.5 * shears_chance[i] + 0.5 * collection_shears_chance for i in range(4)]
P_click_shears = [P_click[i] * P_shears[i] for i in range(4)]

P_click, P_shears, P_click_shears
```

Output:
```
([0.46875, 0.71875, 0.21875, 0.34375],
 [0.125, 0.0625, 0.0625, 0.1875],
 [0.05859375, 0.044921875, 0.013671875, 0.064453125])
```

<!-- grade_id: cell-544c7fce7e22a413, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_click), list)
assert_equal(len(P_click), 4)
```

<!-- grade_id: cell-4d667ae597154cac, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_shears), list)
assert_equal(len(P_shears), 4)
```

<!-- grade_id: cell-bfe5b75eb01e6fdb, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(P_click_shears), list)
assert_equal(len(P_click_shears), 4)
```

## Exercise 12.8
Using the  calculations in Exercise 12.7 as inspiration or as examples where appropriate, write one sentence each  describing  the  treatment  that  the  model  in  Equation (12.10) gives to each of the following quantities. Include whether it is present in the model or not and whether the effect is raw or scaled.

a. Term frequency in a document

b. Collection frequency of a term

c. Document frequency of a term

d. Length normalization of a term

<!-- grade_id: cell-9dda7fcb647eefde, points: 2, grade: true, solution: true -->
a. The term frequency is scaled by the length of the document.

b. This gives a constant chance to documents, scaled with the total number of termen, even if the term doesn't occur in the document.

c. Not used.

d. Not used.

## Exercise 12.9
In the mixture model approach to the query likelihood model (Equation (12.12)), the probability estimate of a term is based on the term frequency of a word in a document, and the collection frequency of the word.  Doing this certainly guarantees that each term of a query (in the vocabulary) has a non-zero chance of being generated by each document. But it has a more subtle but important effect of implementing a form of term weighting, related to what we saw in Chapter 6. Explain how this works. In particular, include  in  your answer a concrete numeric example showing  this  termweighting at work.

<!-- grade_id: cell-b7c0df0f7263b34b, points: 2, grade: true, solution: true -->
1. Als een term niet in het document voorkomt, betekent dat niet direct dat het document een kans van 0 krijgt.
2. Veel voorkomende woorden (stopwoorden) krijgen minder invloed over de uiteindelijke kans, want de impact van het aan- of afwezig zijn van deze woorden is kleiner in verhouding tot de impact van minder voorkomende woorden.

Stel je hebt de volgende dataset:

1. de boss
2. the musician
2. the doctor
3. the master

Dan zou de volgende query: “the boss” toch document 1 als hoogste ranken met de andere documenten op gedeelde 2de plek. Dit terwijl alle documenten voor de helft bestaan uit de helft van de woorden van de query.

$$\begin{align}
P(D_i\mid q)&\propto P(D_i)\prod_{t\in q}(1-\lambda)P(t\mid M_c)-\lambda P(t\mid M_i) \\&\qquad\text{Assuming }P(D_i)\text{ is constant for all documents }i\text{.}\\[5mm]
P(D_1\mid q)&\propto \bigl(0.5P(\texttt{"the"}\mid M_c)+0.5 P(\texttt{"the"}\mid M_1)\bigr) \times \bigl(0.5P(\texttt{"boss"}\mid M_c)+0.5 P(\texttt{"boss"}\mid M_1)\bigr)\\
&\propto 0.5\left(\frac38+\frac02\right) \times 0.5\left(\frac18+\frac12\right)\\
&\propto \frac{15}{256}\\[5mm]
P(D_{j\neq1}\mid q)&\propto \bigl(0.5P(\texttt{"the"}\mid M_c)+0.5 P(\texttt{"the"}\mid M_j)\bigr) \times \bigl(0.5P(\texttt{"boss"}\mid M_c)+0.5 P(\texttt{"boss"}\mid M_j)\bigr)\\
&\propto 0.5\left(\frac38+\frac12\right) \times 0.5\left(\frac18+\frac02\right)\\
&\propto \frac7{256}
\end{align}$$

Dit geeft zowel het effect van 1 als 2 aan.

```python

```