# Zoekmachines

## Notebook made by

__Name__|Erik| en vrienden
__Student id__: (secret)

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Zoekmachines" data-toc-modified-id="Zoekmachines-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Zoekmachines</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li></ul></li><li><span><a href="#Week-5" data-toc-modified-id="Week-5-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Week 5</a></span></li><li><span><a href="#Evaluation-of-Document-Classification-(4pt)" data-toc-modified-id="Evaluation-of-Document-Classification-(4pt)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Evaluation of Document Classification (4pt)</a></span></li><li><span><a href="#Relevance-feedback" data-toc-modified-id="Relevance-feedback-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Relevance feedback</a></span></li><li><span><a href="#Book-excercises-(Chapter-12)" data-toc-modified-id="Book-excercises-(Chapter-12)-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Book excercises (Chapter 12)</a></span><ul class="toc-item"><li><span><a href="#Exercise-12.3" data-toc-modified-id="Exercise-12.3-5.1"><span class="toc-item-num">5.1&nbsp;&nbsp;</span>Exercise 12.3</a></span></li><li><span><a href="#Exercise-12.4" data-toc-modified-id="Exercise-12.4-5.2"><span class="toc-item-num">5.2&nbsp;&nbsp;</span>Exercise 12.4</a></span></li><li><span><a href="#Exercise-12.5" data-toc-modified-id="Exercise-12.5-5.3"><span class="toc-item-num">5.3&nbsp;&nbsp;</span>Exercise 12.5</a></span></li><li><span><a href="#Exercise-12.6" data-toc-modified-id="Exercise-12.6-5.4"><span class="toc-item-num">5.4&nbsp;&nbsp;</span>Exercise 12.6</a></span></li><li><span><a href="#Exercise-12.7" data-toc-modified-id="Exercise-12.7-5.5"><span class="toc-item-num">5.5&nbsp;&nbsp;</span>Exercise 12.7</a></span></li><li><span><a href="#Exercise-12.8" data-toc-modified-id="Exercise-12.8-5.6"><span class="toc-item-num">5.6&nbsp;&nbsp;</span>Exercise 12.8</a></span></li><li><span><a href="#Exercise-12.9" data-toc-modified-id="Exercise-12.9-5.7"><span class="toc-item-num">5.7&nbsp;&nbsp;</span>Exercise 12.9</a></span></li></ul></li></ul></div>

# Week 5

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

# Evaluation of Document Classification (4pt)


Stel je hebt een Naive Bayes Document classifier getrained voor 1 klasse $C$. Stel je hebt een test-collectie van 400 documenten. 

Beantwoord de volgende vragen in de cellen hieronder.

1. De classifier heeft elk document in de collectie geklassificeerd. Wat betekent het als je een precisie hebt van 60%? Hoeveel documenten zijn dan onterecht als behorende tot klasse $C$ geklassificeerd?
2. Stel dat 300 van de test documenten in klasse $C$ vallen? Welke precisie baseline hoort bij deze test-collectie? Welke recall baseline? 
3. Stel je classifier heeft een recall van 50%. Stel 100 van de 400 docs zitten in klasse $C$.  De classifier hoeft niet elk document te classificeren.   Geef exact   aan wat de precisie kan zijn voor deze classifier. Geef dus de onder en bovengrens voor de precisie. 
4. Stel 100 van de 400 docs zitten in klasse $C$. Welke baseline heb je voor _accurracy_?

Met een classificatie baseline bedoelen we een classifier die helemaal niet naar de input kijkt maar alleen slim kiest op basis van de verdeling van de documenten over de klasse $C$.
Denk bijvoorbeeld aan de vraag naar het geslacht van een willekeurige KI student. Wat zou je kiezen als je niks van die student weet? Wat is de recall en de precisie voor het geslacht "jongen" dan?

<!-- grade_id: dc1-a, points: 0, grade: false, solution: true -->
```python
# Deelvraag 1
# Als je een precisie hebt van 60%, zijn 60% van de documenten die geklassificeerd zijn als
# behorende bij klasse C ook daadwerkelijk correct geklassificeerd.

# De vraag is raar, want we zouden moeten weten hoeveel van de 400 docs als klasse C zijn geklassificeerd.
onterecht = 400 * 0.4
```

<!-- grade_id: dc1-b, points: 0, grade: false, solution: true -->
```python
# Deelvraag 2
# De baseline classifier zou dan altijd voor "wel behorende tot klasse C" kiezen.
precisie_baseline = 300 / 400
recall_baseline = 300 / 300
```

<!-- grade_id: dc1-c, points: 0, grade: false, solution: true -->
```python
# Deelvraag 3
# Een recall van 50% betekent dat de helft van de relevante documenten ook daadwerkelijk
# geretrieved zijn. In dit geval is dat 100 / 2 = 50 documenten.
# Als hij ze allemaal heeft geklassificeerd (worst case), is de precisie 50 / 400.
# Als de classifier maar 50 documenten heeft geklassificeerd (best case), is de precisie 50 / 50.
ondergrens = 50 / 400
bovengrens = 400 / 400
```

<!-- grade_id: dc1-d, points: 0, grade: false, solution: true -->
```python
# Deelvraag 4
# De baseline classifier zou dan altijd voor "niet behorende tot klasse C" kiezen.
# De accuracy is per definitie de fractie van klassificaties die correct zijn.
accuracy_baseline = 300 / 400
```

<!-- grade_id: dc1at, points: 1, grade: true, solution: false -->
```python
assert 0<=onterecht<= 400


```

<!-- grade_id: dc1bt, points: 1, grade: true, solution: false -->
```python
assert 0<=precisie_baseline<=1
assert 0<=recall_baseline<= 1

```

<!-- grade_id: dc1ct, points: 1, grade: true, solution: false -->
```python
assert 0<=ondergrens<=1
assert 0<=bovengrens<= 1
```

<!-- grade_id: dc1dt, points: 1, grade: true, solution: false -->
```python
assert 0<=accuracy_baseline<=1

```

# Relevance feedback

In deze vraag gaan we de SMART versie van Rochio relevance feedback (vergelijking 9.3 in het boek) implementeren in Python. We gaan uit van een collectie met een vocabulair van 100 termen. Elke query en elk document is gerepresenteerd als een vector van lengte 100 gevuld met getallen tussen 0 en 1. 

Maak de functies voor `centroid` en `modified_query` af. Alle data zit in numpy arrays, dus je kunt er heel makkelijk mee werken op een manier die je kent uit linaire algebra. 

De data bestaat uit 
* een query van 2 woorden
* 8 relevante documenten en 3 niet relevante documenten

Hint: beide functies kan je zonder for loops schrijven, zelfs comprehensies zijn niet nodig.

```python
# data
q=np.zeros(100)
q[8],q[47]= .4, .111
Relevant_docs=np.array([np.random.rand(100) for _ in range(8)])
NonRelevant_docs=np.array([np.random.rand(100) for _ in range(3)])

print(q)
```

Output:
```
[0.    0.    0.    0.    0.    0.    0.    0.    0.4   0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.111
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.   ]
```

<!-- grade_id: rfa, points: 0, grade: false, solution: true -->
```python
def centroid(D):
    return np.mean(D, axis=0)

# test
print(centroid(NonRelevant_docs))

def modified_query(query, relevantdocs, nonrelevantdocs, alpha=1, beta=0.75, gamma=0.25):
    return alpha * query + beta * centroid(relevantdocs) - gamma * centroid(nonrelevantdocs)

# test
modified_query(q, Relevant_docs, NonRelevant_docs)
```

Output:
```
[0.17909179 0.59335264 0.83842981 0.74371179 0.46907318 0.38976199
 0.45651783 0.57921546 0.63793527 0.53113566 0.62821467 0.28354834
 0.6303849  0.61989842 0.52768373 0.74837886 0.5563328  0.78499047
 0.64221177 0.67422889 0.5478278  0.34260046 0.08571276 0.81795858
 0.75265195 0.60681078 0.45024239 0.63386612 0.50097363 0.39648221
 0.45844656 0.22619958 0.31513494 0.32546177 0.3597536  0.4789071
 0.67472516 0.37525132 0.54202819 0.66699278 0.61987847 0.3648352
 0.60430073 0.72296926 0.26732482 0.41347825 0.57580596 0.36096223
 0.59647746 0.49970271 0.32455147 0.37621593 0.68492863 0.46156835
 0.67216782 0.23396653 0.18281912 0.58910427 0.53094182 0.74980059
 0.33803835 0.49251192 0.28817973 0.72870666 0.40123388 0.44532621
 0.34408423 0.55932974 0.35892648 0.55864328 0.64400226 0.73622889
 0.65689297 0.79923584 0.50395235 0.6617185  0.61849747 0.45294075
 0.50653679 0.46091238 0.35751015 0.47346095 0.69867445 0.27618448
 0.61097203 0.54258209 0.46311815 0.43651826 0.42452158 0.690281
```

<!-- grade_id: rft, points: 1, grade: true, solution: false -->
```python
assert isinstance(modified_query(q,Relevant_docs,NonRelevant_docs), np.ndarray)

```

# Book excercises (Chapter 12)
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

## Exercise 12.3
What  is  the  likelihood  ratio  of  the  document  according  to M1 and M2 in  Example 12.2?

<!-- grade_id: cell-1078c13a222d6f56, points: 0, grade: false, solution: true -->
```python
likelihood_ratio = 0.00000000000048 / 0.000000000000000384
```

<!-- grade_id: cell-d5b4ca4da32d2e9c, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_ratio) in [int, float]
```

## Exercise 12.4
No explicit STOP probability appeared in Example 12.2. Assuming that the STOP probability of each model is 0.1, does this change the likelihood ratio of a document according to the two models?

<!-- grade_id: cell-88d84f5f3e274f8c, points: 0, grade: false, solution: true -->
```python
# It does not influence the likelihood ratio of a document according to
# the two models, since the document has a fixed length and thus causes
# the likelihoods to be multiplied by 0.9 and 0.1 an equal amount of times.
likelihood_with_stop = likelihood_ratio
likelihood_with_stop
```

Output:
```
1250.0
```

<!-- grade_id: cell-77b507102f5e246e, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_with_stop) in [int, float]
```

## Exercise 12.5
How might a language model be used in a spelling correction system? In particular, consider the case of context-sensitive spelling correction, and correcting incorrect usages of words, such as _their_ in _Are you their?_(See Section3.5(page65) for pointers to some literature on this topic.)

<!-- grade_id: cell-4bf9772cefb54700, points: 1, grade: true, solution: true -->
```python
# One idea is to use a bigram language model, where we use the 
# probability of a term occurring after another term. In the case of 
# "Are you their", we would see that the probability of "their" 
# occurring after "you" is very small, and we could correct it by 
# choosing a term that is (1) more likely to occur after "you" and (2)
# similar to the original term. Via this method, we hopefully land on 
# the term "there".
```

## Exercise 12.6
Consider making a language model from the following training text:
    the martian has landed on the latin pop sensation ricky martin
    
a. Under a MLE-estimated unigram probability model, what are P(the) and P(martian)?

b. Under a MLE-estimated bigram model, what are P(sensation|pop) and P(pop|the)?

<!-- grade_id: cell-ecc32bc65a9fbd20, points: 1, grade: false, solution: true -->
```python
def unigram_LM(d: list, t: str):
    """The probability of a term in a document in a unigram LM."""
    return d.count(t) / len(d)

def digram_LM(d: list, t: str, given_t: str):
    """The probability of a term in a document in a digram LM."""
    return sum(d[i] == given_t and d[i + 1] == t
               for i in range(len(d) - 1))

d = "the martian has landed on the latin pop sensation ricky martin".split()
P_the = unigram_LM(d, "the")
P_martian = unigram_LM(d, "martian")
P_sensation_pop = digram_LM(d, "sensation", "pop")
P_pop_the = digram_LM(d, "pop", "the")

P_the, P_martian, P_sensation_pop, P_pop_the
```

Output:
```
(0.18181818181818182, 0.09090909090909091, 1, 0)
```

<!-- grade_id: cell-4bf550a6e30c9cce, points: 0.25, grade: true, solution: false -->
```python
assert type(P_the) in [int, float]
```

<!-- grade_id: cell-cf1055abc4db1355, points: 0.25, grade: true, solution: false -->
```python
assert type(P_martian) in [int, float]
```

<!-- grade_id: cell-c7730c5c9cd22dc6, points: 0.25, grade: true, solution: false -->
```python
assert type(P_sensation_pop) in [int, float]
```

<!-- grade_id: cell-e1b617addf2268f7, points: 0.25, grade: true, solution: false -->
```python
assert type(P_pop_the) in [int, float]
```

## Exercise 12.7
Suppose we have  a collection that consists of the 4 documents given in the below table.

| docID | Document text                              |
|-------|--------------------------------------------|
| 1     | click go the shears boys click click click |
| 2     | click click                                |
| 3     | metal here                                 |
| 4     | metal shears click here                    |


Build a  query likelihood language  model for  this  document  collection. Assume  amixture model between the documents and the collection, with both weighted at 0.5. Maximum likelihood estimation (mle) is used to estimate both as unigram models. Work out the model probabilities of the queries `click`, `shears`, and hence `click shears` for each document, and use those probabilities to rank the documents returned by eachquery.

For each query give the probablities in a list, where the first value corresponds to Doc1, the second value to Doc2 etc.

<!-- grade_id: cell-5a45af37333b7b02, points: 0, grade: false, solution: true -->
```python
def linear_interp_LM_term_prob(d: list, t: str, c: list, gamma: float):
    """The term probability in a document given by a linear interpolation LM."""
    return gamma * unigram_LM(d, t) + (1 - gamma) * unigram_LM(list(np.ravel(c)), t)

def linear_interp_LM_query_likelihood(d: list, q: str, c: list, gamma: float):
    """The likelihood of a query being generated by a document."""
    return np.prod([linear_interp_LM_term_prob(d, t, c, gamma) for t in q.split()])

c = ["click go the shears boys click click click".split(),
     "click click".split(),
     "metal here".split(),
     "metal shears click here".split()]

P_click = [linear_interp_LM_query_likelihood(d, "click", c, 0.5) for d in c]
P_shears = [linear_interp_LM_query_likelihood(d, "shears", c, 0.5) for d in c]
P_click_shears = [linear_interp_LM_query_likelihood(d, "click shears", c, 0.5) for d in c]

P_click, P_shears, P_click_shears
```

Output:
```
/home/erik/.local/lib/python3.8/site-packages/numpy/core/_asarray.py:136: VisibleDeprecationWarning: Creating an ndarray from ragged nested sequences (which is a list-or-tuple of lists-or-tuples-or ndarrays with different lengths or shapes) is deprecated. If you meant to do this, you must specify 'dtype=object' when creating the ndarray
  return array(a, dtype, copy=False, order=order, subok=True)
([0.25, 0.5, 0.0, 0.125],
 [0.0625, 0.0, 0.0, 0.125],
 [0.015625, 0.0, 0.0, 0.015625])
```

<!-- grade_id: cell-544c7fce7e22a413, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_click), list)
assert_equal(len(P_click), 4)
```

<!-- grade_id: cell-4d667ae597154cac, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_shears), list)
assert_equal(len(P_shears), 4)
```

<!-- grade_id: cell-bfe5b75eb01e6fdb, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(P_click_shears), list)
assert_equal(len(P_click_shears), 4)
```

## Exercise 12.8
Using the  calculations in Exercise 12.7 as inspiration or as examples where appropriate, write one sentence each  describing  the  treatment  that  the  model  in  Equation (12.10) gives to each of the following quantities. Include whether it is present in the model or not and whether the effect is raw or scaled.

a. Term frequency in a document

b. Collection frequency of a term

c. Document frequency of a term

d. Length normalization of a term

<!-- grade_id: cell-9dda7fcb647eefde, points: 2, grade: true, solution: true -->
a. The higher gamma is, the more this quantity is scaled up.

b. The lower gamma is, the more this quantity is scaled up.

c. This quantity isn't used in the model.

d. This quantity isn't used in the model.

## Exercise 12.9
In the mixture model approach to the query likelihood model (Equation (12.12)), the probability estimate of a term is based on the term frequency of a word in a document, and the collection frequency of the word.  Doing this certainly guarantees that each term of a query (in the vocabulary) has a non-zero chance of being generated by each document. But it has a more subtle but important effect of implementing a form of term weighting, related to what we saw in Chapter 6. Explain how this works. In particular, include  in  your answer a concrete numeric example showing  this  termweighting at work.

<!-- grade_id: cell-b7c0df0f7263b34b, points: 2, grade: true, solution: true -->
$\gamma$ = 1 causes the query to be interpreted as a conjunction of terms. 
$\gamma$ = 0 on the other hand causes the query to be interpreted as a disjunction of terms.
In general, the higher $\gamma$ is, the "more conjunctive" the query will be interpreted. This means that individual terms are weighted more and get can get more "veto power".

Concrete numeric example using the collection in exercise 12.7, with our example document $d$ being "click click" and our example query $q$ being "click shears":
$$
\text{For } \gamma = 0.7: \\
\begin{align*}
    P(q | d) &= \left(0.7 \cdot \frac{2}{2} + (1 - 0.7) \cdot \frac{8}{16}\right) \cdot \left(0.7 \cdot \frac{0}{2} + (1 - 0.7) \cdot \frac{2}{16}\right) \\
    &= 0.85 \cdot 0.0375 \\
    &= 0.031875
\end{align*} \\
$$
$$
\text{For } \gamma = 0.3: \\
\begin{align*}
    P(q | d) &= \left(0.3 \cdot \frac{2}{2} + (1 - 0.3) \cdot \frac{8}{16}\right) \cdot \left(0.3 \cdot \frac{0}{2} + (1 - 0.3) \cdot \frac{2}{16}\right) \\
    &= 0.65 \cdot 0.0875 \\
    &= 0.056875
\end{align*}
$$

This example makes it clear that for a higher value of $\gamma$, the term "shears" is weighted less. This confirms our theory, since this term doesn't occur in $d$ at all and thus has more "veto power" if $\gamma$ is higher.

```python

```