# Zoekmachines

## Notebook made by

__Name__|Groep| en vrienden

__Student id__ : 

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Zoekmachines" data-toc-modified-id="Zoekmachines-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Zoekmachines</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li></ul></li><li><span><a href="#Week-5" data-toc-modified-id="Week-5-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Week 5</a></span></li><li><span><a href="#Evaluation-of-Document-Classification-(4pt)" data-toc-modified-id="Evaluation-of-Document-Classification-(4pt)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Evaluation of Document Classification (4pt)</a></span></li><li><span><a href="#Relevance-feedback" data-toc-modified-id="Relevance-feedback-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Relevance feedback</a></span></li><li><span><a href="#Book-excercises-(Chapter-12)" data-toc-modified-id="Book-excercises-(Chapter-12)-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Book excercises (Chapter 12)</a></span><ul class="toc-item"><li><span><a href="#Exercise-12.3" data-toc-modified-id="Exercise-12.3-5.1"><span class="toc-item-num">5.1&nbsp;&nbsp;</span>Exercise 12.3</a></span></li><li><span><a href="#Exercise-12.4" data-toc-modified-id="Exercise-12.4-5.2"><span class="toc-item-num">5.2&nbsp;&nbsp;</span>Exercise 12.4</a></span></li><li><span><a href="#Exercise-12.5" data-toc-modified-id="Exercise-12.5-5.3"><span class="toc-item-num">5.3&nbsp;&nbsp;</span>Exercise 12.5</a></span></li><li><span><a href="#Exercise-12.6" data-toc-modified-id="Exercise-12.6-5.4"><span class="toc-item-num">5.4&nbsp;&nbsp;</span>Exercise 12.6</a></span></li><li><span><a href="#Exercise-12.7" data-toc-modified-id="Exercise-12.7-5.5"><span class="toc-item-num">5.5&nbsp;&nbsp;</span>Exercise 12.7</a></span></li><li><span><a href="#Exercise-12.8" data-toc-modified-id="Exercise-12.8-5.6"><span class="toc-item-num">5.6&nbsp;&nbsp;</span>Exercise 12.8</a></span></li><li><span><a href="#Exercise-12.9" data-toc-modified-id="Exercise-12.9-5.7"><span class="toc-item-num">5.7&nbsp;&nbsp;</span>Exercise 12.9</a></span></li></ul></li></ul></div>

# Week 5

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

# Evaluation of Document Classification (4pt)


Stel je hebt een Naive Bayes Document classifier getrained voor 1 klasse $C$. Stel je hebt een test-collectie van 400 documenten. 

Beantwoord de volgende vragen in de cellen hieronder.

1. De classifier heeft elk document in de collectie geklassificeerd. Wat betekent het als je een precisie hebt van 60%? Hoeveel documenten zijn dan onterecht als behorende tot klasse $C$ geklassificeerd?
2. Stel dat 300 van de test documenten in klasse $C$ vallen? Welke precisie baseline hoort bij deze test-collectie? Welke recall baseline? 
3. Stel je classifier heeft een recall van 50%. Stel 100 van de 400 docs zitten in klasse $C$.  De classifier hoeft niet elk document te classificeren.   Geef exact   aan wat de precisie kan zijn voor deze classifier. Geef dus de onder en bovengrens voor de precisie. 
4. Stel 100 van de 400 docs zitten in klasse $C$. Welke baseline heb je voor _accurracy_?

Met een classificatie baseline bedoelen we een classifier die helemaal niet naar de input kijkt maar alleen slim kiest op basis van de verdeling van de documenten over de klasse $C$.
Denk bijvoorbeeld aan de vraag naar het geslacht van een willekeurige KI student. Wat zou je kiezen als je niks van die student weet? Wat is de recall en de precisie voor het geslacht "jongen" dan?

<!-- grade_id: dc1-a, points: 0, grade: false, solution: true -->
```python
# deel vraag 1
onterecht=400*0.4 # replace with your answer
onterecht


```

<!-- grade_id: dc1-b, points: 0, grade: false, solution: true -->
```python
# deel vraag 2

precisie_baseline= 3/4  # replace with your answer
recall_baseline=3/4


```

<!-- grade_id: dc1-c, points: 0, grade: false, solution: true -->
```python
# deel vraag 3

ondergrens= 1/8  # replace with your answer
bovengrens= 1


```

<!-- grade_id: dc1-d, points: 0, grade: false, solution: true -->
```python
# deel vraag 4
# Wat is deze vraag? accuracy is (tp + tn) / (tp + tn + fp + fn) we krijgen hier enkel dat 100 van de 400
# docs in klasse c zitten. Hoe moet je hiervan een baseline accuracy bepalen?

accuracy_baseline= 1  # replace with your answer
 
    

```

<!-- grade_id: dc1at, points: 1, grade: true, solution: false -->
```python
assert 0<=onterecht<= 400


```

<!-- grade_id: dc1bt, points: 1, grade: true, solution: false -->
```python
assert 0<=precisie_baseline<=1
assert 0<=recall_baseline<= 1

```

<!-- grade_id: dc1ct, points: 1, grade: true, solution: false -->
```python
assert 0<=ondergrens<=1
assert 0<=bovengrens<= 1
```

<!-- grade_id: dc1dt, points: 1, grade: true, solution: false -->
```python
assert 0<=accuracy_baseline<=1

```

# Relevance feedback

In deze vraag gaan we de SMART versie van Rochio relevance feedback (vergelijking 9.3 in het boek) implementeren in Python. We gaan uit van een collectie met een vocabulair van 100 termen. Elke query en elk document is gerepresenteerd als een vector van lengte 100 gevuld met getallen tussen 0 en 1. 

Maak de functies voor `centroid` en `modified_query` af. Alle data zit in numpy arrays, dus je kunt er heel makkelijk mee werken op een manier die je kent uit linaire algebra. 

De data bestaat uit 
* een query van 2 woorden
* 8 relevante documenten en 3 niet relevante documenten

Hint: beide functies kan je zonder for loops schrijven, zelfs comprehensies zijn niet nodig.

```python
# data
q=np.zeros(100)
q[8],q[47]= .4, .111
Relevant_docs=np.array([np.random.rand(100) for _ in range(8)])
NonRelevant_docs=np.array([np.random.rand(100) for _ in range(3)])

print(q)
```

<!-- grade_id: rfa, points: 0, grade: false, solution: true -->
```python
def centroid(D):
    return np.mean(D, axis=0)

# test
print(centroid(NonRelevant_docs ))

def modified_query(query,relevantdocs,nonrelevantdocs,alpha=1,beta=.75,gamma=.25):
    len_rel = len(relevantdocs)
    len_non = len(nonrelevantdocs)
    
    a = (alpha * query) + (beta * (1/len_rel) * centroid(relevantdocs)) - (gamma * (1/len_non) * centroid(nonrelevantdocs))

    return a

# test

modified_query(q,Relevant_docs,NonRelevant_docs)


```

<!-- grade_id: rft, points: 1, grade: true, solution: false -->
```python
assert isinstance(modified_query(q,Relevant_docs,NonRelevant_docs), np.ndarray)

```

# Book excercises (Chapter 12)
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

## Exercise 12.3
What  is  the  likelihood  ratio  of  the  document  according  to M1 and M2 in  Example 12.2?

<!-- grade_id: cell-1078c13a222d6f56, points: 0, grade: false, solution: true -->
```python
likelihood_ratio = 0.00000000000048/0.000000000000000384

likelihood_ratio
```

<!-- grade_id: cell-d5b4ca4da32d2e9c, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_ratio) in [int, float]
```

## Exercise 12.4
No explicit STOP probability appeared in Example 12.2. Assuming that the STOP probability of each model is 0.1, does this change the likelihood ratio of a document according to the two models?

<!-- grade_id: cell-88d84f5f3e274f8c, points: 0, grade: false, solution: true -->
```python
likelihood_with_stop = likelihood_ratio
# No adding the word stop does not change the ratio, because it is fixed for both models.

likelihood_with_stop
```

<!-- grade_id: cell-77b507102f5e246e, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_with_stop) in [int, float]
```

## Exercise 12.5
How might a language model be used in a spelling correction system? In particular, consider the case of context-sensitive spelling correction, and correcting incorrect usages of words, such as _their_ in _Are you their?_(See Section3.5(page65) for pointers to some literature on this topic.)

<!-- grade_id: cell-4bf9772cefb54700, points: 1, grade: true, solution: true -->
```python
# One option would be to create an index that guesses what the next word should be based on the words 
# that occured previously.
# Something that could help with this is making a dictionary that contains words that are often (wrongly)
# interchanged with one another. Then when one of these words is used we need to take additional care
# if the given wordchoice fits the context of a given sentence. 
```

## Exercise 12.6
Consider making a language model from the following training text:
    the martian has landed on the latin pop sensation ricky martin
    
a. Under a MLE-estimated unigram probability model, what are P(the) and P(martian)?

b. Under a MLE-estimated bigram model, what are P(sensation|pop) and P(pop|the)?

<!-- grade_id: cell-ecc32bc65a9fbd20, points: 1, grade: false, solution: true -->
```python
text = "the martian has landed on the latin pop sensation ricky martin"
len_text = len(text.split())


P_the = text.count('the') / len_text
P_martian = text.count('martian') / len_text
P_sensation_pop = 1
P_pop_the = 0

P_the, P_martian, P_sensation_pop, P_pop_the
print(P_the, P_martian, P_sensation_pop, P_pop_the)
len(text),text.count("the"), text.count("martian")
```

<!-- grade_id: cell-4bf550a6e30c9cce, points: 0.25, grade: true, solution: false -->
```python
assert type(P_the) in [int, float]
```

<!-- grade_id: cell-cf1055abc4db1355, points: 0.25, grade: true, solution: false -->
```python
assert type(P_martian) in [int, float]
```

<!-- grade_id: cell-c7730c5c9cd22dc6, points: 0.25, grade: true, solution: false -->
```python
assert type(P_sensation_pop) in [int, float]
```

<!-- grade_id: cell-e1b617addf2268f7, points: 0.25, grade: true, solution: false -->
```python
assert type(P_pop_the) in [int, float]
```

## Exercise 12.7
Suppose we have  a collection that consists of the 4 documents given in the below table.

| docID | Document text                              |
|-------|--------------------------------------------|
| 1     | click go the shears boys click click click |
| 2     | click click                                |
| 3     | metal here                                 |
| 4     | metal shears click here                    |


Build a  query likelihood language  model for  this  document  collection. Assume  amixture model between the documents and the collection, with both weighted at 0.5. Maximum likelihood estimation (mle) is used to estimate both as unigram models. Work out the model probabilities of the queries `click`, `shears`, and hence `click shears` for each document, and use those probabilities to rank the documents returned by eachquery.

For each query give the probablities in a list, where the first value corresponds to Doc1, the second value to Doc2 etc.

<!-- grade_id: cell-5a45af37333b7b02, points: 0, grade: false, solution: true -->
```python
doc1 = "click go the shears boys click click click "
doc2 = "click click "
doc3 = "metal here "
doc4 = "metal shears click here"
tot_doc = doc1 + doc2 + doc3 + doc4
lab = 0.5

def Mle(doc, tot_doc, lab, querry):
    l_doc = len(doc.split())
    l_tot = len(tot_doc.split())
    s_querry = querry.split()
    
    if len(s_querry) == 1:
        doc_count = doc.count(querry)
        tot_count = tot_doc.count(querry)
        
        return ((1-lab) * (doc_count / l_doc)) + (lab * (tot_count/l_tot))
    else:
        sum1 = 1 
        for word in s_querry:
            sum1 *= Mle(doc, tot_doc, lab, word)
        
        return sum1


P_click = [Mle(doc1, tot_doc, lab, "click"), Mle(doc2, tot_doc, lab, "click"),
          Mle(doc3, tot_doc, lab, "click"), Mle(doc4, tot_doc, lab, "click")]
P_shears = [Mle(doc1, tot_doc, lab, "shears"), Mle(doc2, tot_doc, lab, "shears"),
          Mle(doc3, tot_doc, lab, "shears"), Mle(doc4, tot_doc, lab, "shears")]
P_click_shears = [Mle(doc1, tot_doc, lab, "click shears"), Mle(doc2, tot_doc, lab, "click shears"),
                 Mle(doc3, tot_doc, lab, "click shears"), Mle(doc4, tot_doc, lab, "click shears")]

P_click, P_shears, P_click_shears
# Ranking doc4 > doc1 > doc2 > doc3
```

<!-- grade_id: cell-544c7fce7e22a413, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_click), list)
assert_equal(len(P_click), 4)
```

<!-- grade_id: cell-4d667ae597154cac, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_shears), list)
assert_equal(len(P_shears), 4)
```

<!-- grade_id: cell-bfe5b75eb01e6fdb, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(P_click_shears), list)
assert_equal(len(P_click_shears), 4)
```

## Exercise 12.8
Using the  calculations in Exercise 12.7 as inspiration or as examples where appropriate, write one sentence each  describing  the  treatment  that  the  model  in  Equation (12.10) gives to each of the following quantities. Include whether it is present in the model or not and whether the effect is raw or scaled.

a. Term frequency in a document

b. Collection frequency of a term

c. Document frequency of a term

d. Length normalization of a term

<!-- grade_id: cell-9dda7fcb647eefde, points: 2, grade: true, solution: true -->
A. Scaled, looking at the word click in the previous excersise we see that it gets a score of .469, if this was not scaled then the value should be 0.5

B. Raw, the overall occurence in the collection itself remains unchanged.

C. Scaled, this is done by the second half of the equation we can see this happening in the querry "shears",even though it doesn't occur in doc 2 the chance given by the equation is not zero.

D. Scaled, once again looking at the previous excersise looking at the term "click" then the second document shows a score of .718 even though it is the only term present in the document.

## Exercise 12.9
In the mixture model approach to the query likelihood model (Equation (12.12)), the probability estimate of a term is based on the term frequency of a word in a document, and the collection frequency of the word.  Doing this certainly guarantees that each term of a query (in the vocabulary) has a non-zero chance of being generated by each document. But it has a more subtle but important effect of implementing a form of term weighting, related to what we saw in Chapter 6. Explain how this works. In particular, include  in  your answer a concrete numeric example showing  this  termweighting at work.

<!-- grade_id: cell-b7c0df0f7263b34b, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```