# Zoekmachines

## Notebook made by

__Name__|Daniël| en vrienden

__Student id__ : secret

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Zoekmachines" data-toc-modified-id="Zoekmachines-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Zoekmachines</a></span><ul class="toc-item"><li><span><a href="#Notebook-made-by" data-toc-modified-id="Notebook-made-by-1.1"><span class="toc-item-num">1.1&nbsp;&nbsp;</span>Notebook made by</a></span></li><li><span><a href="#Toelichting" data-toc-modified-id="Toelichting-1.2"><span class="toc-item-num">1.2&nbsp;&nbsp;</span>Toelichting</a></span></li><li><span><a href="#Voor-het-inleveren!" data-toc-modified-id="Voor-het-inleveren!-1.3"><span class="toc-item-num">1.3&nbsp;&nbsp;</span>Voor het inleveren!</a></span></li></ul></li><li><span><a href="#Week-5" data-toc-modified-id="Week-5-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Week 5</a></span></li><li><span><a href="#Evaluation-of-Document-Classification-(4pt)" data-toc-modified-id="Evaluation-of-Document-Classification-(4pt)-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Evaluation of Document Classification (4pt)</a></span></li><li><span><a href="#Relevance-feedback" data-toc-modified-id="Relevance-feedback-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Relevance feedback</a></span></li><li><span><a href="#Book-excercises-(Chapter-12)" data-toc-modified-id="Book-excercises-(Chapter-12)-5"><span class="toc-item-num">5&nbsp;&nbsp;</span>Book excercises (Chapter 12)</a></span><ul class="toc-item"><li><span><a href="#Exercise-12.3" data-toc-modified-id="Exercise-12.3-5.1"><span class="toc-item-num">5.1&nbsp;&nbsp;</span>Exercise 12.3</a></span></li><li><span><a href="#Exercise-12.4" data-toc-modified-id="Exercise-12.4-5.2"><span class="toc-item-num">5.2&nbsp;&nbsp;</span>Exercise 12.4</a></span></li><li><span><a href="#Exercise-12.5" data-toc-modified-id="Exercise-12.5-5.3"><span class="toc-item-num">5.3&nbsp;&nbsp;</span>Exercise 12.5</a></span></li><li><span><a href="#Exercise-12.6" data-toc-modified-id="Exercise-12.6-5.4"><span class="toc-item-num">5.4&nbsp;&nbsp;</span>Exercise 12.6</a></span></li><li><span><a href="#Exercise-12.7" data-toc-modified-id="Exercise-12.7-5.5"><span class="toc-item-num">5.5&nbsp;&nbsp;</span>Exercise 12.7</a></span></li><li><span><a href="#Exercise-12.8" data-toc-modified-id="Exercise-12.8-5.6"><span class="toc-item-num">5.6&nbsp;&nbsp;</span>Exercise 12.8</a></span></li><li><span><a href="#Exercise-12.9" data-toc-modified-id="Exercise-12.9-5.7"><span class="toc-item-num">5.7&nbsp;&nbsp;</span>Exercise 12.9</a></span></li></ul></li></ul></div>

# Week 5

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

# Evaluation of Document Classification (4pt)


Stel je hebt een Naive Bayes Document classifier getrained voor 1 klasse $C$. Stel je hebt een test-collectie van 400 documenten. 

Beantwoord de volgende vragen in de cellen hieronder.

1. De classifier heeft elk document in de collectie geklassificeerd. Wat betekent het als je een precisie hebt van 60%? Hoeveel documenten zijn dan onterecht als behorende tot klasse $C$ geklassificeerd?
2. Stel dat 300 van de test documenten in klasse $C$ vallen? Welke precisie baseline hoort bij deze test-collectie? Welke recall baseline? 
3. Stel je classifier heeft een recall van 50%. Stel 100 van de 400 docs zitten in klasse $C$.  De classifier hoeft niet elk document te classificeren.   Geef exact   aan wat de precisie kan zijn voor deze classifier. Geef dus de onder en bovengrens voor de precisie. 
4. Stel 100 van de 400 docs zitten in klasse $C$. Welke baseline heb je voor _accurracy_?

Met een classificatie baseline bedoelen we een classifier die helemaal niet naar de input kijkt maar alleen slim kiest op basis van de verdeling van de documenten over de klasse $C$.
Denk bijvoorbeeld aan de vraag naar het geslacht van een willekeurige KI student. Wat zou je kiezen als je niks van die student weet? Wat is de recall en de precisie voor het geslacht "jongen" dan?

<!-- grade_id: dc1-a, points: 0, grade: false, solution: true -->
```python
# deel vraag 1
onterecht= 400 * 0.4 # 60% precision means that 40% must be false
onterecht
```

Output:
```
160.0
```

<!-- grade_id: dc1-b, points: 0, grade: false, solution: true -->
```python
# deel vraag 2
#Precision quantifies the number of positive class predictions that actually belong to the positive class.
#Recall quantifies the number of positive class predictions made out of all positive examples in the dataset.
precisie_baseline= 300/400  # relevant/retrieved
recall_baseline= 300/300 # relevant / total relevant (tp/fn)
```

<!-- grade_id: dc1-c, points: 0, grade: false, solution: true -->
```python
# deel vraag 3

# precision given 50% recall
ondergrens= 50/350  # so in the worst case it wont find any hits
bovengrens=50/50 # at most it can retrieve 50 relevant
```

<!-- grade_id: dc1-d, points: 0, grade: false, solution: true -->
```python
# deel vraag 4

accuracy_baseline= 300/400  #100/400  # accuracy = correct (tP tN) / total
```

<!-- grade_id: dc1at, points: 1, grade: true, solution: false -->
```python
assert 0<=onterecht<= 400


```

<!-- grade_id: dc1bt, points: 1, grade: true, solution: false -->
```python
assert 0<=precisie_baseline<=1
assert 0<=recall_baseline<= 1

```

<!-- grade_id: dc1ct, points: 1, grade: true, solution: false -->
```python
assert 0<=ondergrens<=1
assert 0<=bovengrens<= 1
```

<!-- grade_id: dc1dt, points: 1, grade: true, solution: false -->
```python
assert 0<=accuracy_baseline<=1

```

# Relevance feedback

In deze vraag gaan we de SMART versie van Rochio relevance feedback (vergelijking 9.3 in het boek) implementeren in Python. We gaan uit van een collectie met een vocabulair van 100 termen. Elke query en elk document is gerepresenteerd als een vector van lengte 100 gevuld met getallen tussen 0 en 1. 

Maak de functies voor `centroid` en `modified_query` af. Alle data zit in numpy arrays, dus je kunt er heel makkelijk mee werken op een manier die je kent uit linaire algebra. 

De data bestaat uit 
* een query van 2 woorden
* 8 relevante documenten en 3 niet relevante documenten

Hint: beide functies kan je zonder for loops schrijven, zelfs comprehensies zijn niet nodig.

```python
# data
q=np.zeros(100)
q[8],q[47]= .4, .111
Relevant_docs=np.array([np.random.rand(100) for _ in range(8)])
NonRelevant_docs=np.array([np.random.rand(100) for _ in range(3)])

print(q)
```

Output:
```
[0.    0.    0.    0.    0.    0.    0.    0.    0.4   0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.111
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.    0.
 0.    0.    0.    0.   ]
```

<!-- grade_id: rfa, points: 0, grade: false, solution: true -->
```python
def centroid(D):
    return sum(D) / len(D)

# test
print(centroid(NonRelevant_docs))

def modified_query(query,relevantdocs,nonrelevantdocs,alpha=1,beta=.75,gamma=.25):
    return (alpha * query) + (beta * centroid(relevantdocs)) - (gamma * centroid(nonrelevantdocs))

# test
modified_query(q,Relevant_docs,NonRelevant_docs)
```

Output:
```
[0.70701364 0.53951493 0.67691241 0.36208429 0.48966628 0.49600756
 0.47123047 0.37171822 0.65540407 0.54476505 0.65967553 0.41345226
 0.41169441 0.60940323 0.3616292  0.69655829 0.66517895 0.3445528
 0.25656919 0.54954869 0.36032107 0.68974561 0.57218872 0.20325994
 0.70509773 0.43343794 0.51167537 0.55740563 0.60421628 0.5770204
 0.25122296 0.62429957 0.82031547 0.22937743 0.65154545 0.47997234
 0.5557078  0.25437254 0.51197972 0.59296728 0.55145087 0.32978608
 0.35764872 0.5383642  0.83155337 0.62395072 0.61950942 0.59293685
 0.74392192 0.42163432 0.46368265 0.55438262 0.6372891  0.30055378
 0.56353609 0.58177759 0.45133994 0.56127478 0.44381417 0.39795629
 0.30972583 0.38283968 0.29841306 0.79477909 0.65279104 0.41368092
 0.34434782 0.82013481 0.78431977 0.67897219 0.56606883 0.43936931
 0.52862248 0.20795467 0.50481524 0.35286202 0.57517609 0.25400467
 0.61484631 0.58107526 0.50053059 0.39516101 0.7646507  0.48282601
 0.62886099 0.54269128 0.67261669 0.44083129 0.35028786 0.496363
```

<!-- grade_id: rft, points: 1, grade: true, solution: false -->
```python
assert isinstance(modified_query(q,Relevant_docs,NonRelevant_docs), np.ndarray)

```

# Book excercises (Chapter 12)
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

## Exercise 12.3
What  is  the  likelihood  ratio  of  the  document  according  to M1 and M2 in  Example 12.2?

<!-- grade_id: cell-1078c13a222d6f56, points: 0, grade: false, solution: true -->
```python
m1 = {'the' : .2, 'a' : .1, 'frog' : .01, 'toad' : .01, 'said' : .03,
      'likes' : .02, 'that' : .04, 'dog' : .005, 'cat' : .003, 'monkey' : .001}
m2 = {'the' : .15, 'a' : .12, 'frog' : .0002, 'toad' : .0001, 'said' : .03,
      'likes' : .04, 'that' : .04, 'dog' : .01, 'cat' : .015, 'monkey' : .002}

s = 'frog said that toad likes that dog'.split()

s_given_m1 = np.prod([m1[word] for word in s])
s_given_m2 = np.prod([m2[word] for word in s])

likelihood_ratio = (s_given_m1/s_given_m2).item()

likelihood_ratio
```

Output:
```
1250.0
```

<!-- grade_id: cell-d5b4ca4da32d2e9c, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_ratio) in [int, float]
```

## Exercise 12.4
No explicit STOP probability appeared in Example 12.2. Assuming that the STOP probability of each model is 0.1, does this change the likelihood ratio of a document according to the two models?

<!-- grade_id: cell-88d84f5f3e274f8c, points: 0, grade: false, solution: true -->
```python
# equal stop probability does not alter the likelihood ratio ever.

s_given_m1_stop = np.prod([m1[word] * 0.9 for word in s])
s_given_m2_stop = np.prod([m2[word] * 0.9 for word in s])

likelihood_with_stop = (s_given_m1_stop/ s_given_m2_stop).item()

likelihood_with_stop
```

Output:
```
1249.9999999999998
```

<!-- grade_id: cell-77b507102f5e246e, points: 1, grade: true, solution: false -->
```python
assert type(likelihood_with_stop) in [int, float]
```

## Exercise 12.5
How might a language model be used in a spelling correction system? In particular, consider the case of context-sensitive spelling correction, and correcting incorrect usages of words, such as _their_ in _Are you their?_(See Section3.5(page65) for pointers to some literature on this topic.)

<!-- grade_id: cell-4bf9772cefb54700, points: 1, grade: true, solution: true -->
If a word has a very small probability in a large language model, it is likely spelled wrong. After this the query can be viewed as a probibalistic corruption of a correct query. By finding a word that closely resembles the query but with a higher probability, the system can correct spelling.

## Exercise 12.6
Consider making a language model from the following training text:
    the martian has landed on the latin pop sensation ricky martin
    
a. Under a MLE-estimated unigram probability model, what are P(the) and P(martian)?

b. Under a MLE-estimated bigram model, what are P(sensation|pop) and P(pop|the)?

<!-- grade_id: cell-ecc32bc65a9fbd20, points: 1, grade: false, solution: true -->
```python
text = "the martian has landed on the latin pop sensation ricky martin"

P_the =  text.count("the")/len(text.split())
P_martian = text.count("martian")/len(text.split())
P_sensation_pop = text.count("pop sensation")/text.count("pop")
P_pop_the = text.count("the pop")/text.count("the")

print(P_the, P_martian, P_sensation_pop, P_pop_the)
print(len(text),text.count("the"), text.count("martian"))
```

Output:
```
0.18181818181818182 0.09090909090909091 1.0 0.0
62 2 1
```

<!-- grade_id: cell-4bf550a6e30c9cce, points: 0.25, grade: true, solution: false -->
```python
assert type(P_the) in [int, float]
```

<!-- grade_id: cell-cf1055abc4db1355, points: 0.25, grade: true, solution: false -->
```python
assert type(P_martian) in [int, float]
```

<!-- grade_id: cell-c7730c5c9cd22dc6, points: 0.25, grade: true, solution: false -->
```python
assert type(P_sensation_pop) in [int, float]
```

<!-- grade_id: cell-e1b617addf2268f7, points: 0.25, grade: true, solution: false -->
```python
assert type(P_pop_the) in [int, float]
```

## Exercise 12.7
Suppose we have  a collection that consists of the 4 documents given in the below table.

| docID | Document text                              |
|-------|--------------------------------------------|
| 1     | click go the shears boys click click click |
| 2     | click click                                |
| 3     | metal here                                 |
| 4     | metal shears click here                    |


Build a  query likelihood language  model for  this  document  collection. Assume  amixture model between the documents and the collection, with both weighted at 0.5. Maximum likelihood estimation (mle) is used to estimate both as unigram models. Work out the model probabilities of the queries `click`, `shears`, and hence `click shears` for each document, and use those probabilities to rank the documents returned by eachquery.

For each query give the probablities in a list, where the first value corresponds to Doc1, the second value to Doc2 etc.

<!-- grade_id: cell-5a45af37333b7b02, points: 0, grade: false, solution: true -->
```python
docs = ["click go the shears boys click click click", "click click", "metal here", "metal shears click here"]

click_t = sum([doc.count("click") for doc in docs])
shears_t = sum([doc.count("shears") for doc in docs])
len_t = sum([len(doc) for doc in docs])

P_click = [((doc.count("click")/len(doc.split()) + (click_t/len_t))*0.5) for doc in docs]
P_shears = [((doc.count("shears")/len(doc.split()) + (shears_t/len_t))*0.5) for doc in docs]
P_click_shears = [click * shears for click, shears in zip(P_click, P_shears)]

P_click, P_shears, P_click_shears
```

Output:
```
([0.29069767441860467,
  0.5406976744186046,
  0.040697674418604654,
  0.16569767441860467],
 [0.07412790697674418,
  0.011627906976744186,
  0.011627906976744186,
  0.13662790697674418],
 [0.021548810167658194,
  0.006287182260681449,
  0.00047322877230935646,
  0.022638926446727964])
```

<!-- grade_id: cell-544c7fce7e22a413, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_click), list)
assert_equal(len(P_click), 4)
```

<!-- grade_id: cell-4d667ae597154cac, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(P_shears), list)
assert_equal(len(P_shears), 4)
```

<!-- grade_id: cell-bfe5b75eb01e6fdb, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(P_click_shears), list)
assert_equal(len(P_click_shears), 4)
```

## Exercise 12.8
Using the  calculations in Exercise 12.7 as inspiration or as examples where appropriate, write one sentence each  describing  the  treatment  that  the  model  in  Equation (12.10) gives to each of the following quantities. Include whether it is present in the model or not and whether the effect is raw or scaled.

a. Term frequency in a document

b. Collection frequency of a term

c. Document frequency of a term

d. Length normalization of a term

<!-- grade_id: cell-9dda7fcb647eefde, points: 2, grade: true, solution: true -->
a) yes, scaled (because we divide by lenth)

b) yes, scaled (because we divide by total tokens)

c) no

d) yes, scaled (because we normalise by length of list)

## Exercise 12.9
In the mixture model approach to the query likelihood model (Equation (12.12)), the probability estimate of a term is based on the term frequency of a word in a document, and the collection frequency of the word.  Doing this certainly guarantees that each term of a query (in the vocabulary) has a non-zero chance of being generated by each document. But it has a more subtle but important effect of implementing a form of term weighting, related to what we saw in Chapter 6. Explain how this works. In particular, include  in  your answer a concrete numeric example showing  this  termweighting at work.

<!-- grade_id: cell-b7c0df0f7263b34b, points: 2, grade: true, solution: true -->
So we have a common word that occurs 8 times in d1 and once in d2

And we have a rare word which does not occur in d1 and occurs once in d2

Then P("rare common"|d1)=0.5∗(0/8 + 1/16)∗0.5∗(8/8+9/16)= 0.0244140625

P("rare common"|d2)=0.5∗(1/8 + 1/16)∗0.5∗(1/8 + 9/16)= 0.0322265625

You can see that while d1 has 8 hits and d2 has 2 hits, d2 gets a higher score. The smoothing makes a lot of the same term in a sentence less important, and makes the absence of a term very important.

```python

```