# Zoekmachines

## Notebook made by

__Name__|Erik| en vrienden

__Student id__ : (secret)

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Week-2" data-toc-modified-id="Week-2-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Week 2</a></span></li><li><span><a href="#Book-exercises" data-toc-modified-id="Book-exercises-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Book exercises</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#2.1" data-toc-modified-id="2.1-2.0.1"><span class="toc-item-num">2.0.1&nbsp;&nbsp;</span>2.1</a></span></li><li><span><a href="#2.2-(not-graded)" data-toc-modified-id="2.2-(not-graded)-2.0.2"><span class="toc-item-num">2.0.2&nbsp;&nbsp;</span>2.2 (not graded)</a></span></li><li><span><a href="#2.3" data-toc-modified-id="2.3-2.0.3"><span class="toc-item-num">2.0.3&nbsp;&nbsp;</span>2.3</a></span></li><li><span><a href="#2.4-(not-graded)" data-toc-modified-id="2.4-(not-graded)-2.0.4"><span class="toc-item-num">2.0.4&nbsp;&nbsp;</span>2.4 (not graded)</a></span></li><li><span><a href="#2.8-(not-graded)" data-toc-modified-id="2.8-(not-graded)-2.0.5"><span class="toc-item-num">2.0.5&nbsp;&nbsp;</span>2.8 (not graded)</a></span></li><li><span><a href="#2.9" data-toc-modified-id="2.9-2.0.6"><span class="toc-item-num">2.0.6&nbsp;&nbsp;</span>2.9</a></span></li><li><span><a href="#2.10" data-toc-modified-id="2.10-2.0.7"><span class="toc-item-num">2.0.7&nbsp;&nbsp;</span>2.10</a></span></li><li><span><a href="#6.9" data-toc-modified-id="6.9-2.0.8"><span class="toc-item-num">2.0.8&nbsp;&nbsp;</span>6.9</a></span></li><li><span><a href="#6.10" data-toc-modified-id="6.10-2.0.9"><span class="toc-item-num">2.0.9&nbsp;&nbsp;</span>6.10</a></span></li><li><span><a href="#6.11-(not-graded)" data-toc-modified-id="6.11-(not-graded)-2.0.10"><span class="toc-item-num">2.0.10&nbsp;&nbsp;</span>6.11 (not graded)</a></span></li><li><span><a href="#6.12-(not-graded)" data-toc-modified-id="6.12-(not-graded)-2.0.11"><span class="toc-item-num">2.0.11&nbsp;&nbsp;</span>6.12 (not graded)</a></span></li><li><span><a href="#6.13-(not-graded)" data-toc-modified-id="6.13-(not-graded)-2.0.12"><span class="toc-item-num">2.0.12&nbsp;&nbsp;</span>6.13 (not graded)</a></span></li><li><span><a href="#6.14-(not-graded)" data-toc-modified-id="6.14-(not-graded)-2.0.13"><span class="toc-item-num">2.0.13&nbsp;&nbsp;</span>6.14 (not graded)</a></span></li><li><span><a href="#6.15" data-toc-modified-id="6.15-2.0.14"><span class="toc-item-num">2.0.14&nbsp;&nbsp;</span>6.15</a></span></li><li><span><a href="#6.16-(not-graded)" data-toc-modified-id="6.16-(not-graded)-2.0.15"><span class="toc-item-num">2.0.15&nbsp;&nbsp;</span>6.16 (not graded)</a></span></li><li><span><a href="#6.17" data-toc-modified-id="6.17-2.0.16"><span class="toc-item-num">2.0.16&nbsp;&nbsp;</span>6.17</a></span></li></ul></li></ul></li><li><span><a href="#Extra-theory-exercises" data-toc-modified-id="Extra-theory-exercises-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Extra theory exercises</a></span><ul class="toc-item"><li><span><a href="#Index-size" data-toc-modified-id="Index-size-3.1"><span class="toc-item-num">3.1&nbsp;&nbsp;</span>Index size</a></span></li><li><span><a href="#Index-size-2" data-toc-modified-id="Index-size-2-3.2"><span class="toc-item-num">3.2&nbsp;&nbsp;</span>Index size 2</a></span></li></ul></li><li><span><a href="#Programming-(Inverted-Index)" data-toc-modified-id="Programming-(Inverted-Index)-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Programming (Inverted Index)</a></span><ul class="toc-item"><li><span><a href="#9-Ranked-Retrieval" data-toc-modified-id="9-Ranked-Retrieval-4.1"><span class="toc-item-num">4.1&nbsp;&nbsp;</span>9 Ranked Retrieval</a></span></li><li><span><a href="#9.1-TF-IDF" data-toc-modified-id="9.1-TF-IDF-4.2"><span class="toc-item-num">4.2&nbsp;&nbsp;</span>9.1 TF-IDF</a></span><ul class="toc-item"><li><span><a href="#9.2-TF-IDF-Ranked-Search" data-toc-modified-id="9.2-TF-IDF-Ranked-Search-4.2.1"><span class="toc-item-num">4.2.1&nbsp;&nbsp;</span>9.2 TF-IDF Ranked Search</a></span></li><li><span><a href="#9.3-Vector-Space-Model" data-toc-modified-id="9.3-Vector-Space-Model-4.2.2"><span class="toc-item-num">4.2.2&nbsp;&nbsp;</span>9.3 Vector Space Model</a></span></li><li><span><a href="#9.4-Cosine-Similarity" data-toc-modified-id="9.4-Cosine-Similarity-4.2.3"><span class="toc-item-num">4.2.3&nbsp;&nbsp;</span>9.4 Cosine Similarity</a></span></li><li><span><a href="#9.5-Okapi-BM25" data-toc-modified-id="9.5-Okapi-BM25-4.2.4"><span class="toc-item-num">4.2.4&nbsp;&nbsp;</span>9.5 Okapi BM25</a></span></li></ul></li></ul></li></ul></div>

# Week 2

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

# Book exercises
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

### 2.1

Are the following statements true or false?

a. In a Boolean retrieval system, stemming never lowers precision.

b. In a Boolean retrieval system, stemming never lowers recall.

c. Stemming increases the size of the vocabulary.

d. Stemming should be invoked at indexing time but not while processing a query

<!-- grade_id: cell-9fba209db4797a78, points: 0, grade: false, solution: true -->
```python
a = False  # Je krijgt meer resultaten en niet altijd de juiste resultaten.
b = True  # Er komen alleen maar meer documenten bij, want je zoekt algemener.
c = False  # Sommige woorden gaan naar dezelfde stem.
d = False  # De query moet ook gestemd worden, zodat je het kan vinden in je gestemde index.

a, b, c, d
```

Output:
```
(False, True, False, False)
```

<!-- grade_id: cell-fc1c070acbfa02f6, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(a), bool)
```

<!-- grade_id: cell-323357d08fec7701, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(b), bool)
```

<!-- grade_id: cell-8d2d64004b40ab52, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(c), bool)
```

<!-- grade_id: cell-e3f19c4c556a014a, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(d), bool)
```

### 2.2 (not graded)

Suggest what normalized form should be used for these words (including the word
itself as a possibility):
1. ’Cos
2. Shi’ite
3. cont’d
4. Hawai’i
5. O’Rourke

a) 'Cos: cos (Because the number of Google results for Cos , Cos ,cos and cos are exactly same.)<br/>
b) Shi'ite: shiite (Because on Google, shiite gives more results than shi-ite)<br/>
c) cont'd: contd<br/>
d) Hawai'i: hawaii (Hawai'i, hawai'i, Hawaii)<br/>
e) O'Rourke: orourke (O'Rourke, o'rourke, ORourke)<br/>

### 2.3

The following pairs of words are stemmed to the same form by the Porter stemmer.
Which pairs would you argue shouldn’t be conflated. Give your reasoning.
1. abandon/abandonment
2. absorbency/absorbent
3. marketing/markets
4. university/universe
5. volume/volumes

<!-- grade_id: 203, points: 0, grade: false, solution: true -->
```python
should = set([1, 2, 5])
shouldnot = set([3, 4])

# abandon and abandonment almost mean the same, this also goes for absorbency and absorbent.
# marketing and markets mean two different things, the subject and a place where you can buy stuff.
# university is a educational building, universe is our universe.
# volume can mean the same, but also the loudnes of sound.
```

<!-- grade_id: 203-t, points: 1, grade: true, solution: false -->
```python
assert_equal(type(should), set)

```

### 2.4 (not graded)

For the Porter stemmer rule group shown in (2.1):
1. What is the purpose of including an identity rule such as SS → SS?
2. Applying just this rule group, what will the following words be stemmed to?
   > _circus_ _canaries_ _boss_
3. What rule should be added to correctly stem pony?
4. The stemming for ponies and pony might seem strange. Does it have a deleterious effect on retrieval? Why or why not?

1. The purpose of including an identity such as SS->SS is to prevent the wrong stemming of words like caress using the rule s->nothing.
2. circus->circu, canaries->canari, boss->boss
3. y->i
4. No, because poni is not a separate word with another meaning.

### 2.8 (not graded)

Assume a biword index. Give an example of a document which will be returned
for a query of New York University but is actually a false positive which should not be
returned.

Research on New York, done at York University.

### 2.9

Shown below is a portion of a positional index in the format:
> term: doc1: $\langle$ position1, position2, . . . $\rangle$; doc2: $\langle$ position1, position2, . . .$\rangle$; etc.

angels: 2: $\langle$36,174,252,651$\rangle$; 4: $\langle$12,22,102,432$\rangle$; 7: $\langle$17$\rangle$;<br/>
fools: 2: $\langle$1,17,74,222$\rangle$; 4: $\langle$8,78,108,458$\rangle$; 7: $\langle$3,13,23,193$\rangle$;<br/>
fear: 2: $\langle$87,704,722,901$\rangle$; 4: $\langle$13,43,113,433$\rangle$; 7: $\langle$18,328,528$\rangle$;<br/>
in: 2: $\langle$3,37,76,444,851$\rangle$; 4: $\langle$10,20,110,470,500$\rangle$; 7: $\langle$5,15,25,195$\rangle$;<br/>
rush: 2: $\langle$2,66,194,321,702$\rangle$; 4: $\langle$9,69,149,429,569$\rangle$; 7: $\langle$4,14,404$\rangle$;<br/>
to: 2: $\langle$47,86,234,999$\rangle$; 4: $\langle$14,24,774,944$\rangle$; 7: $\langle$199,319,599,709$\rangle$;<br/>
tread: 2: $\langle$57,94,333$\rangle$; 4: $\langle$15,35,155$\rangle$; 7: $\langle$20,320$\rangle$;<br/>
where: 2: $\langle$67,124,393,1001$\rangle$; 4: $\langle$11,41,101,421,431$\rangle$; 7: $\langle$16,36,736$\rangle$;

Which document(s) if any match each of the following queries, where each expression
within quotes is a phrase query?
1. “fools rush in”
2. “fools rush in” AND “angels fear to tread”

Give your answer as sets, respectively `firstphrase` and `secondphrase`

<!-- grade_id: cell-04b90595b2cf07b8, points: 0, grade: false, solution: true -->
```python
# doc 2: <1, 2, 3>, doc 4: <8, 9, 10>, doc 7: <3, 4, 5>
firstphrase = set([2, 4, 7])
# doc 4: <12, 13, 14, 15>
secondphrase = set([4])
```

<!-- grade_id: cell-32f1ec171801d213, points: 1, grade: true, solution: false -->
```python
assert_equal(type(firstphrase),set)
assert_equal(type(secondphrase),set)
for x in firstphrase:
    assert x in [2,4,7]
for x in secondphrase:
    assert x in [2,4,7]
```

### 2.10

Consider the following fragment of a positional index with the format:

> word: document: $\langle$position, position, . . .$\rangle$; document: $\langle$position, . . .$\rangle$<br />
. . .

>Gates: 1: $\langle$3$\rangle$; 2: $\langle$6$\rangle$; 3: $\langle$2,17$\rangle$; 4: $\langle$1$\rangle$;<br />
IBM: 4: $\langle$3$\rangle$; 7: $\langle$14$\rangle$;<br />
Microsoft: 1: $\langle$1$\rangle$; 2: $\langle$1,21$\rangle$; 3: $\langle$3$\rangle$; 5: $\langle$16,22,51$\rangle$;

The $/k$ operator, word1 $/k$ word2 finds occurrences of word1 within $k$ words of word2 (on either side), where $k$ is a positive integer argument. Thus $k = 1$ demands that word1 be adjacent to word2.
1. Describe the set of documents that satisfy the query Gates $/2$ Microsoft.
2. Describe for each value  $k$ between 1 and 7 the set of documents returned after the   query Gates $/k$ Microsoft.  



<!-- grade_id: cell-31e611774e33399f, points: 0, grade: false, solution: true -->
```python
k1 = set([3])
k2 = set([1, 3])
k3 = set([1, 3])
k4 = set([1, 3])
k5 = set([1, 2, 3])
k6 = set([1, 2, 3])
k7 = set([1, 2, 3])
```

<!-- grade_id: cell-b933f7947f3f78db, points: 1, grade: true, solution: false -->
```python
for k in [k1,k2,k3,k4,k5,k6,k7]:
    assert_equal(type(k),set)
    assert k.issubset({1,2,3,4,5})
```

### 6.9

What is the idf of a term that occurs in every document (variable `idf_every_doc`)? Compare this with the use of stop word lists.

<!-- grade_id: cell-84323557f1b57844, points: 0, grade: false, solution: true -->
```python
idf_every_doc = 0
```

<!-- grade_id: cell-3cbf108756459c24, points: 0.5, grade: true, solution: false -->
```python
assert type(idf_every_doc) in [int, float]
```

### 6.10

Consider the table of term frequencies for 3 documents denoted Doc1, Doc2, Doc3 in Figure 6.9. Compute the tf-idf weights for the terms car, auto, insurance, best, for each document, using the idf values from Figure 6.8.

Note that we use tf-idf weighting **with the raw tf values** in this question.

<!-- grade_id: cell-3753d6674d83ae0e, points: 0, grade: false, solution: true -->
```python
idf_car = 1.65
idf_auto = 2.08
idf_insurance = 1.62
idf_best = 1.5

# for Doc1
tfidf_car1 = idf_car * 27
tfidf_auto1 = idf_auto * 3
tfidf_insurance1 = idf_insurance * 0
tfidf_best1 = idf_best * 14
# for Doc2
tfidf_car2 = idf_car * 4
tfidf_auto2 = idf_auto * 33
tfidf_insurance2 = idf_insurance * 33
tfidf_best2 = idf_best * 0
# for Doc3
tfidf_car3 = idf_car * 24
tfidf_auto3 = idf_auto * 0
tfidf_insurance3 = idf_insurance * 29
tfidf_best3 = idf_best * 17

print("Doc1: [%s, %s, %s, %s]" %(tfidf_car1, tfidf_auto1, tfidf_insurance1, tfidf_best1))
print("Doc2: [%s, %s, %s, %s]" %(tfidf_car2, tfidf_auto2, tfidf_insurance2, tfidf_best2))
print("Doc3: [%s, %s, %s, %s]" %(tfidf_car3, tfidf_auto3, tfidf_insurance3, tfidf_best3))
```

Output:
```
Doc1: [44.55, 6.24, 0.0, 21.0]
Doc2: [6.6, 68.64, 53.46, 0.0]
Doc3: [39.599999999999994, 0.0, 46.980000000000004, 25.5]
```

<!-- grade_id: cell-d465a3ce0a126ca9, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car1), float)
assert_equal(type(tfidf_auto1), float)
assert_equal(type(tfidf_insurance1), float)
assert_equal(type(tfidf_best1), float)
```

<!-- grade_id: cell-f48eb29d8d2e1cfb, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car2), float)
assert_equal(type(tfidf_auto2), float)
assert_equal(type(tfidf_insurance2), float)
assert_equal(type(tfidf_best2), float)
```

<!-- grade_id: cell-fc354ac597577c33, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car3), float)
assert_equal(type(tfidf_auto3), float)
assert_equal(type(tfidf_insurance3), float)
assert_equal(type(tfidf_best3), float)
```

### 6.11 (not graded)

Can the tf-idf weight of a term in a document exceed 1?

Yes, obviously, as seen above.

### 6.12 (not graded)

How does the base of the logarithm in (6.7) affect the score calculation in (6.9)? How does the base of the logarithm affect the relative scores of two documents on a given query?

For a base $x > 1$, the score will become lower if the base becomes higher. For a base $0 < x < 1$, this still applies, but the score will be negative. Therefore, if the base is smaller, the score from (6.9) will be relatively higher if the query contains some term(s) that occur(s) in very few documents.

Since $log_x$ is a monotonically increasing function if $x > 1$ (i.e. $log_x(a) > log_x(b)$ if $a > b$ for every $x > 1$), the order of the scores of two documents for a given query will be preserved by definition. If $0 < x < 1$, $log_x$ will be a monotionically decreasing function and thus the order of the scores will be reversed.

I'm not sure whether this is enough or the question asks for a factor describing how the relative scores change. So just to be sure, this is calculated below. When the base increases with a factor $f$, the score increases with a factor g:
$$\begin{align*}
    \log_{xf}(a) &= g \cdot \log_x(a) \\
    \frac{\log(a)}{\log(fx)} &= g \cdot \frac{\log(a)}{\log(x)} \\
    g &= \frac{\log(a) \cdot \log(x)}{\log(fx) \cdot \log(a)} \\
    g &= \log_{xf}(x)
\end{align*}$$

### 6.13 (not graded)

If the logarithm in (6.7) is computed base 2, suggest a simple approximation to the idf of a term.

When the base is 2, we could look at the number of binary digits that N/DF(t) has.

### 6.14 (not graded)

If we were to stem jealous and jealousy to a common stem before setting up the vector space, detail how the definitions of tf and idf should be modified.

Jealous and jealousy would become one term. Therefore tf will be the sum of the tf's of the aformentioned terms, thus increase. Idf will possibly decrease as more documents can contain the stemmed term (increased df).

### 6.15

""Recall the tf-idf weights computed in Exercise 6.10. Compute the Euclidean normalized document vectors for each of the documents, where each vector has four components, one for each of the four terms.""

Instead of the question above, make the function `euclidean_normalize(vector)` which normalizes a list (which represents a doc).

<!-- grade_id: cell-c09bec4374c36b0f, points: 1, grade: false, solution: true -->
```python
def euclidean_normalize(vector):
    return list(vector / np.linalg.norm(vector))

x = euclidean_normalize([tfidf_car3, tfidf_auto3, tfidf_insurance3, tfidf_best3])
print(x)
# What test shows that your answer is correct?
print(np.sqrt(np.sum(np.array(x)**2)))
```

Output:
```
[0.5952679795256183, 0.0, 0.7062042848008473, 0.38331650196725425]
1.0
```

<!-- grade_id: cell-2cb078e3927ebd52, points: 1, grade: true, solution: false -->
```python
assert_equal(type(euclidean_normalize([tfidf_car3,tfidf_auto3,tfidf_insurance3,tfidf_best3])), list)
assert_equal(len(euclidean_normalize([tfidf_car3,tfidf_auto3,tfidf_insurance3,tfidf_best3])), 4)
```

### 6.16 (not graded)

Verify that the square root of the sum of the squares of the components of each of the document vectors in Exercise 6.15 is 1 (to within rounding error). Why is this the case?

$\texttt{np.sqrt(np.sum(x**2)) == 1.0}$, see previous exercise.
This is the case, because we normalized the vector.

### 6.17

With term weights as computed in Exercise 6.15, rank the three documents by computed score (using the cosine similarity of the document and the query vector) for the query "car insurance", for each of the following cases of term weighting in the query: 

1. The weight of a term is 1 if present in the query, 0 otherwise. (variable `rank1`)
2.  The weight of a term is the idf of that term. (variable `rank2`)



<!-- grade_id: cell-4a6a0261e6fe0975, points: 0, grade: false, solution: true -->
```python
docs = np.array(list(map(euclidean_normalize, [[tfidf_car1, tfidf_auto1, tfidf_insurance1, tfidf_best1],
                                               [tfidf_car2, tfidf_auto2, tfidf_insurance2, tfidf_best2],
                                               [tfidf_car3, tfidf_auto3, tfidf_insurance3, tfidf_best3]])))
vec1 = np.array(euclidean_normalize([1, 0, 1, 0]))
vec2 = np.array(euclidean_normalize([1.65, 0, 1.62, 0]))

rank1 = sorted(zip([1, 2, 3], docs @ vec1), key=lambda x: -x[1])
rank2 = sorted(zip([1, 2, 3], docs @ vec2), key=lambda x: -x[1])

rank1, rank2
```

Output:
```
([(3, 0.9202798636314545), (1, 0.6345355454218854), (2, 0.48673542014951854)],
 [(3, 0.9195214992872569), (1, 0.6403300253625414), (2, 0.48323104205732087)])
```

<!-- grade_id: cell-0409ff7b0d9afc0d, points: 1, grade: true, solution: false -->
```python
assert_equal(type(rank1),list)
```

# Extra theory exercises

## Index size

Suppose you have a collection of documents with the following properties:

* It has `N` documents.
* Each document contains on average `A` many tokens, and `U` many unique tokens.
* The vocabulary size of the collection is `V` (this means, that there are `V` many unique words/tokens in the collection.

In this question we look at **the size of the posting lists**, when we consider different types of indexes.
The size of the posting list is obtained by "tokenizing" the posting list. So we ignore separators like comma, and list symbols, etc. For instance, the number of tokens (or the "length") of posting `{'d2':(2,[4,18]), 'd8':(1,[7])}` , is 7.

Now calculate for each of the following situations the **total number of tokens in the posting lists of the inverted index**.  

**For each situation, you should consider the most space efficient index. This means that the index size will grow with each situation described below.**

Note the we ask for the size of the posting lists. You do not need to add the space needed for the terms (or termids) themselves.

The situations are given in the answer code cell. Give the answer as a formula using the variables declared below.


```python
N=100 # number of docs
V=8000 # number of unique words in collection
A=300  # avg number of words per doc
U=100  # avg number of unique words per doc 
```

<!-- grade_id: isa, points: 0, grade: false, solution: true -->
```python
# size for an index which is suitable for Boolean search (without any form of ranking)
BS_index_size = N * U
print(BS_index_size)

# size for an index in which we store for each term  the term-frequency in each document
TF_index_size = 2 * N * U
print(TF_index_size)
```

Output:
```
10000
20000
```

<!-- grade_id: ist, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(BS_index_size,int)

# Hier staat een onzichtbare test
```

<!-- grade_id: ist-TF, points: 0.5, grade: true, solution: false -->
```python

assert isinstance(TF_index_size,int)
# Hier staat een onzichtbare test
```

## Index size 2

We continue with the numbers from the previous question and add the following.

1. 40% of all unique words in the vocabulary occur only once in the whole collection (so once in one document). These are called _hapaxes_.
2. On the other hand  2% of all  unique words in the vocabulary occur in at least half of all documents. We call these terms _stopwords_.

Now consider the simplest index from the previous question, which is suitable for Boolean Search.

1.  How many tokens are there in the posting lists of the hapaxes? (`Hapax_size`)
2.  Minimally how many tokens are there in the posting lists of the stopwords?

<!-- grade_id: is2a, points: 0, grade: false, solution: true -->
```python
# How many tokens are there in the posting lists of the hapaxes?
Hapax_size = 0.4 * V
print(Hapax_size)

# Minimally how many tokens are there in the posting lists of the stopwords?
Stopword_size = 0.02 * V * 0.5 * N
print(Stopword_size)
```

Output:
```
3200.0
8000.0
```

<!-- grade_id: is2t, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(Hapax_size,float)

```

<!-- grade_id: is2tbis, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(Stopword_size,float)

```

# Programming (Inverted Index)

```python
from collections import OrderedDict,Counter,defaultdict
from bs4 import BeautifulSoup
import sys
import itertools
import os
import glob
from tqdm import tqdm_notebook
import nltk
import zipfile
import math
import numpy as np
def loadShakespeare():
    if 'shaks200.zip' in os.listdir():
        return 'shaks200.zip'
    elif os.path.exists('../../data/Week1/'):
        return '../../data/Week1/shaks200.zip'
    elif os.path.exists('../../../data/Week1/'):
        return '../../../data/Week1/shaks200.zip'
```

```python
zipfile.ZipFile('../../data/Week1/shaks200.zip')
```

Output:
```
/tmp/ipykernel_7931/2216002993.py in <module>
----> 1 zipfile.ZipFile('../../data/Week1/shaks200.zip')

/usr/lib/python3.8/zipfile.py in __init__(self, file, mode, compression, allowZip64, compresslevel, strict_timestamps)
   1249             while True:
   1250                 try:
-> 1251                     self.fp = io.open(file, filemode)
   1252                 except OSError:
   1253                     if filemode in modeDict:

FileNotFoundError: [Errno 2] No such file or directory: '../../data/Week1/shaks200.zip'
```

<!-- grade_id: cell-512011c182747803, points: 0, grade: false, solution: true -->
```python
docs = []

def index_collection(zipfolder):
    global docs
    # With zipfile we can read the file without opening the zip file
    archive = zipfile.ZipFile(zipfolder, 'r')
    namelist = [x for x in archive.namelist() if '.xml' in x]
    docs = [file[:-4] for file in namelist] 
    index = defaultdict(lambda: defaultdict(int))
                        
    for infile in tqdm_notebook(namelist):  # loop over each file
        f = archive.open(infile)
        filename = infile[:-4]
        
        soup = BeautifulSoup(f.read(), 'html.parser')
        for token in nltk.tokenize.word_tokenize(soup.get_text()):
            if not token or not token.isalpha():
                continue

            index[token.lower()][filename] += 1

    return index

Shakespeare = index_collection(loadShakespeare())
```

Output:
```
/tmp/ipykernel_7931/3888869661.py:11: TqdmDeprecationWarning: This function will be removed in tqdm==5.0.0
Please use `tqdm.notebook.tqdm` instead of `tqdm.tqdm_notebook`
  for infile in tqdm_notebook(namelist):  # loop over each file
  0%|          | 0/37 [00:00<?, ?it/s]
```

<!-- grade_id: cell-512011c182747803a, points: 0, grade: false, solution: true -->
```python
"""Copy your answers to question 3 from week 1 here, 
So NewIndex will be created"""
CorpusFreq = {}
DocFreq = {} 
NewIndex = {}

for w in Shakespeare:
    DocFreq[w] = len(Shakespeare[w])
    CorpusFreq[w] = sum(Shakespeare[w].values())
    NewIndex[w] = {'CorpusFreq': CorpusFreq[w], 'DocFreq': DocFreq[w], 'Freq_per_Doc': Shakespeare[w]}
```

<!-- grade_id: cell-512011c182747803b, points: 0, grade: false, solution: true -->
```python
"""Copy your answers to question 3 from week 1 here,
So NewIndex["TotalFiles"] will be created"""

NewIndex["TotalFiles"] = len(docs)
```

```python
# test the schema of NewIndex. For each term, these attributes should be present.
assert_equal(set(NewIndex['the'].keys()), {'Freq_per_Doc', 'CorpusFreq', 'DocFreq'})
```

<!-- grade_id: cell-0a422d512c7a3d0f, points: 0, grade: true, solution: false -->
```python
assert len(NewIndex) > 1
assert_equal(type(NewIndex["TotalFiles"]), int)
```

## 9 Ranked Retrieval

In this section we extend the boolean search engine to perform more sophisticated ranked retrieval.

## 9.1 TF-IDF
Implement the TF-IDF function 6.8 from the IR -book. So, given a query term $t$ and a document $d$, compute and return the tf-idf$_{t,d}$ score.

Note that we use the raw term frequency here, as in the book. Not the log normalized tf as in the slides.

**Use logarithm with base 2 in this and all further exercises.**

<!-- grade_id: cell-0a466ded8aed5c5c, points: 0, grade: false, solution: true -->
```python
def tfidf(term, document, index):
    """
    Given a single query term and a document, score the document using
    tf-idf.
    """
    return np.log2(index["TotalFiles"] / index[term]["DocFreq"]) \
        * index[term]["Freq_per_Doc"][document]
    
tfidf('brutus', 'a_and_c', NewIndex)
```

Output:
```
8.837813462515799
```

<!-- grade_id: cell-e35bc40f7ab55f2f, points: 1, grade: true, solution: false -->
```python
assert isinstance(tfidf('brutus', 'a_and_c', NewIndex), float)
```

### 9.2 TF-IDF Ranked Search

Implement the scoring function 6.9 from the IR -book. So, given a query $Q$ consisting of a list of words, compute for each document the score for this query, and return the document ids and scores in decreasing order.

Break ties alphabettically on the doc id (so if docs a and b have the same score a should appear before b).

The result is a list of pairs ('docid','score').

<!-- grade_id: cell-966e28b929230aab, points: 0, grade: false, solution: true -->
```python
def rankedSearchTfidf(query, index):
    """
    Given a query, score the documents according to the tf-idf scoring
    scheme and return  the documents and scores in decreasing order.
    Break ties alphabettically on the doc id (so if docs a and b have the same score a should appear before b)
    """
    return sorted([(doc, sum(tfidf(term, doc, index) for term in query)) 
                   for doc in docs],
                  key=lambda d: (-d[1], d[0]))

rankedSearchTfidf(['brutus', 'caesar'], NewIndex)
```

Output:
```
[('j_caesar', 1108.7171966824922),
 ('a_and_c', 286.71878474408607),
 ('coriolan', 240.8304168535555),
 ('cymbelin', 17.30746533933656),
 ('hamlet', 4.132505069999678),
 ('hen_vi_2', 4.132505069999678),
 ('rich_iii', 3.8461034087414574),
 ('hen_v', 3.170979217814314),
 ('titus', 3.170979217814314),
 ('merchant', 2.2094533656289497),
 ('hen_vi_1', 1.9230517043707287),
 ('hen_vi_3', 1.9230517043707287),
 ('m_for_m', 1.9230517043707287),
 ('all_well', 0.9615258521853643),
 ('as_you', 0.9615258521853643),
 ('hen_iv_2', 0.9615258521853643),
 ('lll', 0.9615258521853643),
 ('m_wives', 0.9615258521853643),
 ('macbeth', 0.9615258521853643),
 ('othello', 0.9615258521853643),
 ('rich_ii', 0.9615258521853643),
 ('com_err', 0.0),
 ('dream', 0.0),
 ('hen_iv_1', 0.0),
 ('hen_viii', 0.0),
 ('john', 0.0),
 ('lear', 0.0),
 ('much_ado', 0.0),
 ('pericles', 0.0),
 ('r_and_j', 0.0),
 ('t_night', 0.0),
 ('taming', 0.0),
 ('tempest', 0.0),
 ('timon', 0.0),
 ('troilus', 0.0),
 ('two_gent', 0.0),
 ('win_tale',
```

<!-- grade_id: cell-1ca47df9d70de0ed, points: 1, grade: true, solution: false -->
```python
result = rankedSearchTfidf(['brutus','caesar'], NewIndex)
assert_equal(type(result),list)
assert_equal(type(result[0]),tuple)
assert_equal(type(result[0][0]),str)
assert isinstance(result[0][1],float)
```

### 9.3 Vector Space Model

We now want to implement the cosine similarity function. To do this we need to vectorize documents and normalize them. We use Euclidean length for that (see section 6.3). Create a function that can vectorize and normalize a document:

$normalizedVector(d) = \frac{V(d)}{|V(d)|}$

* Make sure you don't take `TotalFiles` as a word
* Include all words in the corpus, even if they are not in document (so they will be 0)
* Make sure you sort the words alphabetically
* the value for each term is its TF-IDF weight as calculated by your earlier `tfidf` function.

<!-- grade_id: cell-c9694a0ad4a895c6, points: 0, grade: false, solution: true -->
```python
def normalizedVector(document, index):
    """
    Given a document, vectorize and normalize it
    """
    return euclidean_normalize([tfidf(term, document, index)
                                for term in sorted(index.keys())
                                if term != "TotalFiles"])

len(normalizedVector('all_well', NewIndex)), normalizedVector('all_well', NewIndex)[:5]

# Check that the Euclidean length of your vector equals 1
```

Output:
```
(22332, [0.0, 0.0, 0.0, 0.0, 0.0028840422804802812])
```

<!-- grade_id: cell-2e8497964ac87a83, points: 1, grade: true, solution: false -->
```python
assert_equal(len(normalizedVector('all_well', NewIndex)), len(NewIndex)-1)

```

### 9.4 Cosine Similarity

Use the normalized lengths to compute the cosine-similarity between a query and a document. I.e. implement scoring formula 6.12 from the book.

Do this in the form of the function `cosineSimilarity(query, document, index)`

The query vector is a boolean vector with the value for a term being 1 if and only if it occurs at least once in the query.

<!-- grade_id: cell-a87cb124f00cdf2e, points: 0, grade: false, solution: true -->
```python
def cosineSimilarity(query, document, index):
    """
    Given a query and a document, compute and return the cosine-similarity
    between the two.
    """
    q = euclidean_normalize([term in query
                             for term in sorted(index.keys())
                             if term != "TotalFiles"])
    d = normalizedVector(document, index)
    # d and q are already normalized.
    return np.dot(q, d)

cosineSimilarity(['brutus','love'], 'a_and_c', NewIndex)
```

Output:
```
0.0029940455322842877
```

<!-- grade_id: cell-c960058b3d6d3183, points: 1, grade: true, solution: false -->
```python
assert type(cosineSimilarity(['love','lightly'], 'all_well', NewIndex)) in [float, np.float64, np.float32]
```

### 9.5 Okapi BM25
Now create a search engine which scores according to another famous scoring formula [Okapi BM25](https://en.wikipedia.org/wiki/Okapi_BM25).

<!-- grade_id: cell-236a53b720ae81a9, points: 0, grade: false, solution: true -->
```python
def rankedSearchBM25(query, index, a=1.5, b=0.75):
    """
    Given a query, score the documents according to the BM25 scoring
    scheme and return a list of (document, score) pairs  in decreasing order
    """
    def idf(term):
        return np.log2(index["TotalFiles"] / index[term]["DocFreq"])

    def tf(term, doc):
        return index[term]["Freq_per_Doc"][doc]

    doc_lens = {doc: sum(index[term]["Freq_per_Doc"][doc]
                         for term in index.keys()
                         if term != "TotalFiles")
                for doc in docs}
    avgdl = sum(doc_lens.values()) / index["TotalFiles"]

    return sorted([(doc,
                    sum(idf(term) * tf(term, doc) * (a + 1)
                        / (tf(term, doc) + a * (1 - b + b * doc_lens[doc] / avgdl))
                        for term in query)
                   ) for doc in docs],
                  key=lambda d: (-d[1], d[0]))

rankedSearchBM25(['brutus'], NewIndex)
```

Output:
```
[('j_caesar', 5.504143111026695),
 ('coriolan', 5.43654956062449),
 ('a_and_c', 3.917758754230206),
 ('titus', 2.302373094325634),
 ('merchant', 2.2830245460259255),
 ('hen_vi_2', 2.094720547214092),
 ('hen_v', 2.0714511243347946),
 ('hamlet', 1.9156463192454167),
 ('all_well', 0.0),
 ('as_you', 0.0),
 ('com_err', 0.0),
 ('cymbelin', 0.0),
 ('dream', 0.0),
 ('hen_iv_1', 0.0),
 ('hen_iv_2', 0.0),
 ('hen_vi_1', 0.0),
 ('hen_vi_3', 0.0),
 ('hen_viii', 0.0),
 ('john', 0.0),
 ('lear', 0.0),
 ('lll', 0.0),
 ('m_for_m', 0.0),
 ('m_wives', 0.0),
 ('macbeth', 0.0),
 ('much_ado', 0.0),
 ('othello', 0.0),
 ('pericles', 0.0),
 ('r_and_j', 0.0),
 ('rich_ii', 0.0),
 ('rich_iii', 0.0),
 ('t_night', 0.0),
 ('taming', 0.0),
 ('tempest', 0.0),
 ('timon', 0.0),
 ('troilus', 0.0),
 ('two_gent', 0.0),
 ('win_tale', 0.0)]
```

<!-- grade_id: cell-860fbb8c8bd77a2a, points: 1, grade: true, solution: false -->
```python
result = rankedSearchBM25(['brutus'], NewIndex)
assert_equal(type(result),list)
assert_equal(type(result[0]),tuple)
assert_equal(type(result[0][0]),str)
assert isinstance(result[0][1],float)
```

```python

```