# Zoekmachines

## Notebook made by

__Name__|Levi| en vrienden

__Student id__ : secret

__Name__|Raoul| en vrienden

__Student id__ : secret

__Name__|Daniël| en vrienden

__Student id__ : secret

__Name__|Philip| en vrienden

__Student id__ : secret

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar zichtbare tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. De andere tests worden na inleveren toegevoegd aan die cell.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Dit levert problemen op bij nakijken. Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Week-2" data-toc-modified-id="Week-2-1"><span class="toc-item-num">1&nbsp;&nbsp;</span>Week 2</a></span></li><li><span><a href="#Book-exercises" data-toc-modified-id="Book-exercises-2"><span class="toc-item-num">2&nbsp;&nbsp;</span>Book exercises</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#2.1" data-toc-modified-id="2.1-2.0.1"><span class="toc-item-num">2.0.1&nbsp;&nbsp;</span>2.1</a></span></li><li><span><a href="#2.2-(not-graded)" data-toc-modified-id="2.2-(not-graded)-2.0.2"><span class="toc-item-num">2.0.2&nbsp;&nbsp;</span>2.2 (not graded)</a></span></li><li><span><a href="#2.3" data-toc-modified-id="2.3-2.0.3"><span class="toc-item-num">2.0.3&nbsp;&nbsp;</span>2.3</a></span></li><li><span><a href="#2.4-(not-graded)" data-toc-modified-id="2.4-(not-graded)-2.0.4"><span class="toc-item-num">2.0.4&nbsp;&nbsp;</span>2.4 (not graded)</a></span></li><li><span><a href="#2.8-(not-graded)" data-toc-modified-id="2.8-(not-graded)-2.0.5"><span class="toc-item-num">2.0.5&nbsp;&nbsp;</span>2.8 (not graded)</a></span></li><li><span><a href="#2.9" data-toc-modified-id="2.9-2.0.6"><span class="toc-item-num">2.0.6&nbsp;&nbsp;</span>2.9</a></span></li><li><span><a href="#2.10" data-toc-modified-id="2.10-2.0.7"><span class="toc-item-num">2.0.7&nbsp;&nbsp;</span>2.10</a></span></li><li><span><a href="#6.9" data-toc-modified-id="6.9-2.0.8"><span class="toc-item-num">2.0.8&nbsp;&nbsp;</span>6.9</a></span></li><li><span><a href="#6.10" data-toc-modified-id="6.10-2.0.9"><span class="toc-item-num">2.0.9&nbsp;&nbsp;</span>6.10</a></span></li><li><span><a href="#6.11-(not-graded)" data-toc-modified-id="6.11-(not-graded)-2.0.10"><span class="toc-item-num">2.0.10&nbsp;&nbsp;</span>6.11 (not graded)</a></span></li><li><span><a href="#6.12-(not-graded)" data-toc-modified-id="6.12-(not-graded)-2.0.11"><span class="toc-item-num">2.0.11&nbsp;&nbsp;</span>6.12 (not graded)</a></span></li><li><span><a href="#6.13-(not-graded)" data-toc-modified-id="6.13-(not-graded)-2.0.12"><span class="toc-item-num">2.0.12&nbsp;&nbsp;</span>6.13 (not graded)</a></span></li><li><span><a href="#6.14-(not-graded)" data-toc-modified-id="6.14-(not-graded)-2.0.13"><span class="toc-item-num">2.0.13&nbsp;&nbsp;</span>6.14 (not graded)</a></span></li><li><span><a href="#6.15" data-toc-modified-id="6.15-2.0.14"><span class="toc-item-num">2.0.14&nbsp;&nbsp;</span>6.15</a></span></li><li><span><a href="#6.16-(not-graded)" data-toc-modified-id="6.16-(not-graded)-2.0.15"><span class="toc-item-num">2.0.15&nbsp;&nbsp;</span>6.16 (not graded)</a></span></li><li><span><a href="#6.17" data-toc-modified-id="6.17-2.0.16"><span class="toc-item-num">2.0.16&nbsp;&nbsp;</span>6.17</a></span></li></ul></li></ul></li><li><span><a href="#Extra-theory-exercises" data-toc-modified-id="Extra-theory-exercises-3"><span class="toc-item-num">3&nbsp;&nbsp;</span>Extra theory exercises</a></span><ul class="toc-item"><li><span><a href="#Index-size" data-toc-modified-id="Index-size-3.1"><span class="toc-item-num">3.1&nbsp;&nbsp;</span>Index size</a></span></li><li><span><a href="#Index-size-2" data-toc-modified-id="Index-size-2-3.2"><span class="toc-item-num">3.2&nbsp;&nbsp;</span>Index size 2</a></span></li></ul></li><li><span><a href="#Programming-(Inverted-Index)" data-toc-modified-id="Programming-(Inverted-Index)-4"><span class="toc-item-num">4&nbsp;&nbsp;</span>Programming (Inverted Index)</a></span><ul class="toc-item"><li><span><a href="#9-Ranked-Retrieval" data-toc-modified-id="9-Ranked-Retrieval-4.1"><span class="toc-item-num">4.1&nbsp;&nbsp;</span>9 Ranked Retrieval</a></span></li><li><span><a href="#9.1-TF-IDF" data-toc-modified-id="9.1-TF-IDF-4.2"><span class="toc-item-num">4.2&nbsp;&nbsp;</span>9.1 TF-IDF</a></span><ul class="toc-item"><li><span><a href="#9.2-TF-IDF-Ranked-Search" data-toc-modified-id="9.2-TF-IDF-Ranked-Search-4.2.1"><span class="toc-item-num">4.2.1&nbsp;&nbsp;</span>9.2 TF-IDF Ranked Search</a></span></li><li><span><a href="#9.3-Vector-Space-Model" data-toc-modified-id="9.3-Vector-Space-Model-4.2.2"><span class="toc-item-num">4.2.2&nbsp;&nbsp;</span>9.3 Vector Space Model</a></span></li><li><span><a href="#9.4-Cosine-Similarity" data-toc-modified-id="9.4-Cosine-Similarity-4.2.3"><span class="toc-item-num">4.2.3&nbsp;&nbsp;</span>9.4 Cosine Similarity</a></span></li><li><span><a href="#9.5-Okapi-BM25" data-toc-modified-id="9.5-Okapi-BM25-4.2.4"><span class="toc-item-num">4.2.4&nbsp;&nbsp;</span>9.5 Okapi BM25</a></span></li></ul></li></ul></li></ul></div>

# Week 2

```python
from nose.tools import assert_equal, assert_almost_equal
import numpy as np
```

Output:
```
Collecting nose
  Downloading nose-1.3.7-py3-none-any.whl (154 kB)
[?25l[K     |██▏                             | 10 kB 24.0 MB/s eta 0:00:01[K     |████▎                           | 20 kB 27.1 MB/s eta 0:00:01[K     |██████▍                         | 30 kB 12.5 MB/s eta 0:00:01[K     |████████▌                       | 40 kB 9.8 MB/s eta 0:00:01[K     |██████████▋                     | 51 kB 5.1 MB/s eta 0:00:01[K     |████████████▊                   | 61 kB 5.5 MB/s eta 0:00:01[K     |██████████████▉                 | 71 kB 5.9 MB/s eta 0:00:01[K     |█████████████████               | 81 kB 6.6 MB/s eta 0:00:01[K     |███████████████████             | 92 kB 6.7 MB/s eta 0:00:01[K     |█████████████████████▏          | 102 kB 5.2 MB/s eta 0:00:01[K     |███████████████████████▎        | 112 kB 5.2 MB/s eta 0:00:01[K     |█████████████████████████▍      | 122 kB 5.2 MB/s eta 0:00:01[K     |███████████████████████████▌    | 133 kB 5.2 MB/s eta 0:00:01[K     |
```

# Book exercises
In this section you will make the exercises from the book, which is freely available online as a PDF at [nlp.stanford.edu/IR-book/](http://nlp.stanford.edu/IR-book/).

### 2.1

Are the following statements true or false?

a. In a Boolean retrieval system, stemming never lowers precision.

b. In a Boolean retrieval system, stemming never lowers recall.

c. Stemming increases the size of the vocabulary.

d. Stemming should be invoked at indexing time but not while processing a query

<!-- grade_id: cell-9fba209db4797a78, points: 0, grade: false, solution: true -->
```python
a = False # Replace with your answer
b =  True # Replace with your answer
c = False # Replace with your answer
d = False # Replace with your answer

#Ik dacht dus het volgende maar hebben we het tijdens wg wel over - Raoul

# a = False # Replace with your answer
# b = True # Replace with your answer
# c = False # Replace with your answer
# d = True # Replace with your answer


a, b, c, d
```

Output:
```
(False, True, False, False)
```

<!-- grade_id: cell-fc1c070acbfa02f6, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(a), bool)
```

<!-- grade_id: cell-323357d08fec7701, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(b), bool)
```

<!-- grade_id: cell-8d2d64004b40ab52, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(c), bool)
```

<!-- grade_id: cell-e3f19c4c556a014a, points: 0.25, grade: true, solution: false -->
```python
assert_equal(type(d), bool)
```

### 2.2 (not graded)

Suggest what normalized form should be used for these words (including the word
itself as a possibility):
1. ’Cos
2. Shi’ite
3. cont’d
4. Hawai’i
5. O’Rourke

a) 'Cos: cos ( The ' is not a information in text it more a conection between words.)<br/>
b) Shi'ite: shiite (this the same as next)<br/>
c) cont'd: contd<br/>
d) Hawai'i: hawaii (Hawai'i, hawai'i, Hawaii)<br/>
e) O'Rourke: orourke (O'Rourke, o'rourke, ORourke)<br/>

### 2.3

The following pairs of words are stemmed to the same form by the Porter stemmer.
Which pairs would you argue shouldn’t be conflated. Give your reasoning.
1. abandon/abandonment
2. absorbency/absorbent
3. marketing/markets
4. university/universe
5. volume/volumes

<!-- grade_id: 203, points: 0, grade: false, solution: true -->
```python
should=set({1,2,5}) # answer as a set of numbers between 1 and 5
shouldnot=set({3,4})




```

<!-- grade_id: 203-t, points: 1, grade: true, solution: false -->
```python
assert_equal(type(should), set)

```

### 2.4 (not graded)

For the Porter stemmer rule group shown in (2.1):
1. What is the purpose of including an identity rule such as SS → SS?
2. Applying just this rule group, what will the following words be stemmed to?
   > _circus_ _canaries_ _boss_
3. What rule should be added to correctly stem pony?
4. The stemming for ponies and pony might seem strange. Does it have a deleterious effect on retrieval? Why or why not?

1. The purpose of including an identity such as SS->SS is to prevent the wrong stemming of words like caress using the rule s->nothing.
* circus->circu, canaries->canari, boss->boss
* y->i
*  so ponies -> poni and pony -> poni we can see  that poni might get more forms will be stemed to this from. like if ponzi scheme was spelld incorrectly in one of your documents you can get the same word. and this makes it less usefull.


### 2.8 (not graded)

Assume a biword index. Give an example of a document which will be returned
for a query of New York University but is actually a false positive which should not be
returned.

Research on New York, done at York University.

### 2.9

Shown below is a portion of a positional index in the format:
> term: doc1: $\langle$ position1, position2, . . . $\rangle$; doc2: $\langle$ position1, position2, . . .$\rangle$; etc.

angels: 2: $\langle$36,174,252,651$\rangle$; 4: $\langle$12,22,102,432$\rangle$; 7: $\langle$17$\rangle$;<br/>
fools: 2: $\langle$1,17,74,222$\rangle$; 4: $\langle$8,78,108,458$\rangle$; 7: $\langle$3,13,23,193$\rangle$;<br/>
fear: 2: $\langle$87,704,722,901$\rangle$; 4: $\langle$13,43,113,433$\rangle$; 7: $\langle$18,328,528$\rangle$;<br/>
in: 2: $\langle$3,37,76,444,851$\rangle$; 4: $\langle$10,20,110,470,500$\rangle$; 7: $\langle$5,15,25,195$\rangle$;<br/>
rush: 2: $\langle$2,66,194,321,702$\rangle$; 4: $\langle$9,69,149,429,569$\rangle$; 7: $\langle$4,14,404$\rangle$;<br/>
to: 2: $\langle$47,86,234,999$\rangle$; 4: $\langle$14,24,774,944$\rangle$; 7: $\langle$199,319,599,709$\rangle$;<br/>
tread: 2: $\langle$57,94,333$\rangle$; 4: $\langle$15,35,155$\rangle$; 7: $\langle$20,320$\rangle$;<br/>
where: 2: $\langle$67,124,393,1001$\rangle$; 4: $\langle$11,41,101,421,431$\rangle$; 7: $\langle$16,36,736$\rangle$;

Which document(s) if any match each of the following queries, where each expression
within quotes is a phrase query?
1. “fools rush in”
2. “fools rush in” AND “angels fear to tread”

Give your answer as sets, respectively `firstphrase` and `secondphrase`

<!-- grade_id: cell-04b90595b2cf07b8, points: 0, grade: false, solution: true -->
```python
firstphrase = set([2,4,7])
secondphrase = set([4])

```

<!-- grade_id: cell-32f1ec171801d213, points: 1, grade: true, solution: false -->
```python
assert_equal(type(firstphrase),set)
assert_equal(type(secondphrase),set)
for x in firstphrase:
    assert x in [2,4,7]
for x in secondphrase:
    assert x in [2,4,7]
```

### 2.10

Consider the following fragment of a positional index with the format:

> word: document: $\langle$position, position, . . .$\rangle$; document: $\langle$position, . . .$\rangle$<br />
. . .

>Gates: 1: $\langle$3$\rangle$; 2: $\langle$6$\rangle$; 3: $\langle$2,17$\rangle$; 4: $\langle$1$\rangle$;<br />
IBM: 4: $\langle$3$\rangle$; 7: $\langle$14$\rangle$;<br />
Microsoft: 1: $\langle$1$\rangle$; 2: $\langle$1,21$\rangle$; 3: $\langle$3$\rangle$; 5: $\langle$16,22,51$\rangle$;

The $/k$ operator, word1 $/k$ word2 finds occurrences of word1 within $k$ words of word2 (on either side), where $k$ is a positive integer argument. Thus $k = 1$ demands that word1 be adjacent to word2.
1. Describe the set of documents that satisfy the query Gates $/2$ Microsoft.
2. Describe for each value  $k$ between 1 and 7 the set of documents returned after the   query Gates $/k$ Microsoft.  



<!-- grade_id: cell-31e611774e33399f, points: 0, grade: false, solution: true -->
```python
k1 = set([3])
k2 = set([1,3])
k3 = set([1,3])
k4 = set([1,3])
k5 = set([1,2,3])
k6 = set([1,2,3])
k7 = set([1,2,3])

```

<!-- grade_id: cell-b933f7947f3f78db, points: 1, grade: true, solution: false -->
```python
for k in [k1,k2,k3,k4,k5,k6,k7]:
    assert_equal(type(k),set)
    assert k.issubset({1,2,3,4,5})
```

### 6.9

What is the idf of a term that occurs in every document (variable `idf_every_doc`)? Compare this with the use of stop word lists.

<!-- grade_id: cell-84323557f1b57844, points: 0, grade: false, solution: true -->
```python
idf_every_doc = 0

```

<!-- grade_id: cell-3cbf108756459c24, points: 0.5, grade: true, solution: false -->
```python
assert type(idf_every_doc) in [int, float]
```

### 6.10

Consider the table of term frequencies for 3 documents denoted Doc1, Doc2, Doc3 in Figure 6.9. Compute the tf-idf weights for the terms car, auto, insurance, best, for each document, using the idf values from Figure 6.8.

Note that we use tf-idf weighting **with the raw tf values** in this question.

<!-- grade_id: cell-3753d6674d83ae0e, points: 0, grade: false, solution: true -->
```python
# for Doc1
tfidf_car1 = 1.65 * 27
tfidf_auto1 = 2.08 * 3
tfidf_insurance1 = 1.62 * 0
tfidf_best1 =  1.5 * 14
# for Doc2
tfidf_car2 = 1.65 * 4
tfidf_auto2 = 2.08 * 33
tfidf_insurance2 = 1.62 * 33
tfidf_best2 =  1.5 * 0
# for Doc3
tfidf_car3 = 1.65 * 24
tfidf_auto3 = 2.08 * 0
tfidf_insurance3 = 1.62 * 29
tfidf_best3 =  1.5 * 17


print("Doc1: [%s, %s, %s, %s]" %(tfidf_car1, tfidf_auto1, tfidf_insurance1, tfidf_best1))
print("Doc2: [%s, %s, %s, %s]" %(tfidf_car2, tfidf_auto2, tfidf_insurance2, tfidf_best2))
print("Doc3: [%s, %s, %s, %s]" %(tfidf_car3, tfidf_auto3, tfidf_insurance3, tfidf_best3))
```

Output:
```
Doc1: [44.55, 6.24, 0.0, 21.0]
Doc2: [6.6, 68.64, 53.46, 0.0]
Doc3: [39.599999999999994, 0.0, 46.980000000000004, 25.5]
```

<!-- grade_id: cell-d465a3ce0a126ca9, points: 0.34, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car1), float)
assert_equal(type(tfidf_auto1), float)
assert_equal(type(tfidf_insurance1), float)
assert_equal(type(tfidf_best1), float)
```

<!-- grade_id: cell-f48eb29d8d2e1cfb, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car2), float)
assert_equal(type(tfidf_auto2), float)
assert_equal(type(tfidf_insurance2), float)
assert_equal(type(tfidf_best2), float)
```

<!-- grade_id: cell-fc354ac597577c33, points: 0.33, grade: true, solution: false -->
```python
assert_equal(type(tfidf_car3), float)
assert_equal(type(tfidf_auto3), float)
assert_equal(type(tfidf_insurance3), float)
assert_equal(type(tfidf_best3), float)
```

### 6.11 (not graded)

Can the tf-idf weight of a term in a document exceed 1?

yes it can if the if the term has a high tf like a document with one word let say we have a corpus with 8 documents and the word world is in 2 of them this means we have idf of $\log(\frac{8}{2}) = log(4) = 2$ so we have a IDF of 2 and let see if we have a document like " Hello World world " now we have tf = $\frac{2}{3}$ now we have a tf-idf = $\frac{2}{3} * 2 = \frac{4}{3}$

### 6.12 (not graded)

How does the base of the logarithm in (6.7) affect the score calculation in (6.9)? How does the base of the logarithm affect the relative scores of two documents on a given query?

how higher the log is the closer the numbers will be cause the times a number has to multiply by before it reaches the net big break point 

### 6.13 (not graded)

If the logarithm in (6.7) is computed base 2, suggest a simple approximation to the idf of a term.

Computers use base 2 (binary), so by looking at the number of bits that the number needs to be represented, we can make a simple approximation of the idf of a term.

### 6.14 (not graded)

If we were to stem jealous and jealousy to a common stem before setting up the vector space, detail how the definitions of tf and idf should be modified.

the TF would be $tf_{jealous} + tf_{jealousy} = TF_{stemmed \ jealous}$ but with IDF we can't do the same cause there is a set of documents that both contain jealous and jealousy

### 6.15

""Recall the tf-idf weights computed in Exercise 6.10. Compute the Euclidean normalized document vectors for each of the documents, where each vector has four components, one for each of the four terms.""

Instead of the question above, make the function `euclidean_normalize(vector)` which normalizes a list (which represents a doc).

<!-- grade_id: cell-c09bec4374c36b0f, points: 1, grade: false, solution: true -->
```python
def euclidean_normalize(vector):
  return list(vector/np.linalg.norm(vector))
# What test shows that your answer is correct?

```

<!-- grade_id: cell-2cb078e3927ebd52, points: 1, grade: true, solution: false -->
```python
assert_equal(type(euclidean_normalize([tfidf_car3,tfidf_auto3,tfidf_insurance3,tfidf_best3])), list)
assert_equal(len(euclidean_normalize([tfidf_car3,tfidf_auto3,tfidf_insurance3,tfidf_best3])), 4)
```

### 6.16 (not graded)

Verify that the square root of the sum of the squares of the components of each of the document vectors in Exercise 6.15 is 1 (to within rounding error). Why is this the case?

cause the length of the vector is the $\sqrt{n_1^2 + n_1^2 + ... + n_x^2}$ so this the length and when we part that menas if you 

### 6.17

With term weights as computed in Exercise 6.15, rank the three documents by computed score (using the cosine similarity of the document and the query vector) for the query "car insurance", for each of the following cases of term weighting in the query: 

1. The weight of a term is 1 if present in the query, 0 otherwise. (variable `rank1`)
2.  The weight of a term is the idf of that term. (variable `rank2`)



<!-- grade_id: cell-4a6a0261e6fe0975, points: 0, grade: false, solution: true -->
```python
rank1=[1,0,1,0]  # list of numbers from 1,2,3
rank2=[1.65,0,1,62,0]

 

rank1,rank2
```

Output:
```
([1, 0, 1, 0], [1.65, 0, 1, 62, 0])
```

<!-- grade_id: cell-0409ff7b0d9afc0d, points: 1, grade: true, solution: false -->
```python
assert_equal(type(rank1),list)
```

# Extra theory exercises

## Index size

Suppose you have a collection of documents with the following properties:

* It has `N` documents.
* Each document contains on average `A` many tokens, and `U` many unique tokens.
* The vocabulary size of the collection is `V` (this means, that there are `V` many unique words/tokens in the collection.

In this question we look at **the size of the posting lists**, when we consider different types of indexes.
The size of the posting list is obtained by "tokenizing" the posting list. So we ignore separators like comma, and list symbols, etc. For instance, the number of tokens (or the "length") of posting `{'d2':(2,[4,18]), 'd8':(1,[7])}` , is 7.

Now calculate for each of the following situations the **total number of tokens in the posting lists of the inverted index**.  

**For each situation, you should consider the most space efficient index. This means that the index size will grow with each situation described below.**

Note the we ask for the size of the posting lists. You do not need to add the space needed for the terms (or termids) themselves.

The situations are given in the answer code cell. Give the answer as a formula using the variables declared below.


```python
N=100 # number of docs
V=8000 # number of unique words in collection
A=300  # avg number of words per doc
U=100  # avg number of unique words per doc 
```

<!-- grade_id: isa, points: 0, grade: false, solution: true -->
```python
# size for an index which is suitable for Boolean search (without any form of ranking)
BS_index_size = V + U * N 


print(BS_index_size)

# size for an index in which we store for each term  the term-frequency in each document
TF_index_size = V + 2 * U * N 


print(TF_index_size)
```

Output:
```
18000
28000
```

<!-- grade_id: ist, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(BS_index_size,int)

# Hier staat een onzichtbare test
```

<!-- grade_id: ist-TF, points: 0.5, grade: true, solution: false -->
```python

assert isinstance(TF_index_size,int)
# Hier staat een onzichtbare test
```

## Index size 2

We continue with the numbers from the previous question and add the following.

1. 40% of all unique words in the vocabulary occur only once in the whole collection (so once in one document). These are called _hapaxes_.
2. On the other hand  2% of all  unique words in the vocabulary occur in at least half of all documents. We call these terms _stopwords_.

Now consider the simplest index from the previous question, which is suitable for Boolean Search.

1.  How many tokens are there in the posting lists of the hapaxes? (`Hapax_size`)
2.  Minimally how many tokens are there in the posting lists of the stopwords?

<!-- grade_id: is2a, points: 0, grade: false, solution: true -->
```python
# How many tokens are there in the posting lists of the hapaxes?

Hapax_size = V * 0.4


print(Hapax_size)

# Minimally how many tokens are there in the posting lists of the stopwords?
Stopword_size = V * 0.02


print(Stopword_size)
```

Output:
```
3200.0
160.0
```

<!-- grade_id: is2t, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(Hapax_size,float)

```

<!-- grade_id: is2tbis, points: 0.5, grade: true, solution: false -->
```python
assert isinstance(Stopword_size,float)

```

# Programming (Inverted Index)

```python
from collections import OrderedDict,Counter,defaultdict
from bs4 import BeautifulSoup
import sys
import itertools
import os
import glob
from tqdm import tqdm_notebook
import nltk
import zipfile
import math
import numpy as np
def loadShakespeare():
    if 'shaks200.zip' in os.listdir():
        return 'shaks200.zip'
    elif os.path.exists('../../data/Week1/'):
        return '../../data/Week1/shaks200.zip'
    elif os.path.exists('../../../data/Week1/'):
        return '../../../data/Week1/shaks200.zip'

nltk.download('punkt')
```

Output:
```
[nltk_data] Downloading package punkt to /root/nltk_data...
[nltk_data]   Unzipping tokenizers/punkt.zip.
True
```

```python
zipfile.ZipFile('../../data/Week1/shaks200.zip')
```

Output:
```
<ipython-input-18-801c60c80d22> in <module>()
----> 1 zipfile.ZipFile('../../data/Week1/shaks200.zip')

/usr/lib/python3.7/zipfile.py in __init__(self, file, mode, compression, allowZip64, compresslevel)
   1238             while True:
   1239                 try:
-> 1240                     self.fp = io.open(file, filemode)
   1241                 except OSError:
   1242                     if filemode in modeDict:

FileNotFoundError: [Errno 2] No such file or directory: '../../data/Week1/shaks200.zip'
```

<!-- grade_id: cell-512011c182747803, points: 0, grade: false, solution: true -->
```python
"""Copy your answers to question 1 from week 1 here"""
def index_collection(zipfolder):
    # With zipfile we can read the file without opening the zip file
    archive = zipfile.ZipFile(zipfolder, 'r')
    namelist = [x for x in archive.namelist() if '.xml' in x]
    MyIndex = defaultdict(Counter) # initialize MyIndex
    for infile in tqdm_notebook(namelist): # loop over each file
        f = archive.open(infile)
        # Iterate over every word of every doc and add it to the dict
        for word in nltk.word_tokenize(BeautifulSoup(f, "xml").get_text()):
            if word.isalpha():
                MyIndex[word.lower()][f.name[:-4]] += 1
        
    return MyIndex
Shakespeare = index_collection(loadShakespeare())
```

Output:
```
/usr/local/lib/python3.7/dist-packages/ipykernel_launcher.py:7: TqdmDeprecationWarning: This function will be removed in tqdm==5.0.0
Please use `tqdm.notebook.tqdm` instead of `tqdm.tqdm_notebook`
  import sys
  0%|          | 0/37 [00:00<?, ?it/s]
```

<!-- grade_id: cell-512011c182747803a, points: 0, grade: false, solution: true -->
```python
"""Copy your answers to question 3 from week 1 here,
So NewIndex will be created"""
CorpusFreq = {key : sum(Shakespeare[str(key)].values()) for key in Shakespeare.keys()}
DocFreq = {key : len(Shakespeare[str(key)]) for key in Shakespeare.keys()}
NewIndex = {key : {'CorpusFreq' : CorpusFreq[key], 'DocFreq' : DocFreq[key],
             'Freq_per_Doc': Shakespeare[key]} for key in Shakespeare.keys()}

```

<!-- grade_id: cell-512011c182747803b, points: 0, grade: false, solution: true -->
```python
"""Copy your answers to question 3 from week 1 here,
So NewIndex["TotalFiles"] will be created"""

NewIndex["TotalFiles"] = set()
for value in list(NewIndex.values())[:-1]:
    for doc_id in list(value['Freq_per_Doc'].keys()):
        NewIndex["TotalFiles"].add(doc_id)
        
NewIndex["TotalFiles"] = len(NewIndex["TotalFiles"])

```

```python
# test the schema of NewIndex. For each term, these attributes should be present.
assert_equal(set(NewIndex['the'].keys()), {'Freq_per_Doc', 'CorpusFreq', 'DocFreq'})
```

<!-- grade_id: cell-0a422d512c7a3d0f, points: 0, grade: true, solution: false -->
```python
assert len(NewIndex) > 1
assert_equal(type(NewIndex["TotalFiles"]), int)
```

## 9 Ranked Retrieval

In this section we extend the boolean search engine to perform more sophisticated ranked retrieval.

## 9.1 TF-IDF
Implement the TF-IDF function 6.8 from the IR -book. So, given a query term $t$ and a document $d$, compute and return the tf-idf$_{t,d}$ score.

Note that we use the raw term frequency here, as in the book. Not the log normalized tf as in the slides.

**Use logarithm with base 2 in this and all further exercises.**

<!-- grade_id: cell-0a466ded8aed5c5c, points: 0, grade: false, solution: true -->
```python
def tfidf(term, document, index):
    """
    Given a single query term and a document, score the document using
    tf-idf.
    """
    if not  index[term]['Freq_per_Doc'][document]:
      return 0
    tf = index[term]['Freq_per_Doc'][document]
    idf = math.log(index['TotalFiles'] / index[term]['DocFreq'], 2)
    return tf * idf
    
tfidf('brutus', 'a_and_c', NewIndex)

```

Output:
```
8.8378134625158
```

<!-- grade_id: cell-e35bc40f7ab55f2f, points: 1, grade: true, solution: false -->
```python
assert isinstance(tfidf('brutus', 'a_and_c', NewIndex), float)
```

### 9.2 TF-IDF Ranked Search

Implement the scoring function 6.9 from the IR -book. So, given a query $Q$ consisting of a list of words, compute for each document the score for this query, and return the document ids and scores in decreasing order.

Break ties alphabettically on the doc id (so if docs a and b have the same score a should appear before b).

The result is a list of pairs ('docid','score').

<!-- grade_id: cell-966e28b929230aab, points: 0, grade: false, solution: true -->
```python
def rankedSearchTfidf(query,index):
    """
    Given a query, score the documents according to the tf-idf scoring
    scheme and return  the documents and scores in decreasing order.
    Break ties alphabettically on the doc id (so if docs a and b have the same score a should appear before b)
    """
    out = defaultdict(float)
    for term in query:
        for document in index[term]['Freq_per_Doc']:
            out[document] = out[document] + tfidf(term,document,index)
    out = list(out.items())
    out.sort(key=lambda x:x[1], reverse=True)
    return out
    
rankedSearchTfidf(['brutus', 'caesar'], NewIndex)
```

Output:
```
[('j_caesar', 1108.7171966824924),
 ('a_and_c', 286.71878474408607),
 ('coriolan', 240.83041685355556),
 ('cymbelin', 17.30746533933656),
 ('hamlet', 4.132505069999679),
 ('hen_vi_2', 4.132505069999679),
 ('rich_iii', 3.8461034087414574),
 ('hen_v', 3.1709792178143146),
 ('titus', 3.1709792178143146),
 ('merchant', 2.20945336562895),
 ('hen_vi_1', 1.9230517043707287),
 ('hen_vi_3', 1.9230517043707287),
 ('m_for_m', 1.9230517043707287),
 ('all_well', 0.9615258521853643),
 ('as_you', 0.9615258521853643),
 ('hen_iv_2', 0.9615258521853643),
 ('lll', 0.9615258521853643),
 ('m_wives', 0.9615258521853643),
 ('macbeth', 0.9615258521853643),
 ('othello', 0.9615258521853643),
 ('rich_ii', 0.9615258521853643)]
```

<!-- grade_id: cell-1ca47df9d70de0ed, points: 1, grade: true, solution: false -->
```python
result = rankedSearchTfidf(['brutus','caesar'], NewIndex)
assert_equal(type(result),list)
assert_equal(type(result[0]),tuple)
assert_equal(type(result[0][0]),str)
assert isinstance(result[0][1],float)
```

### 9.3 Vector Space Model

We now want to implement the cosine similarity function. To do this we need to vectorize documents and normalize them. We use Euclidean length for that (see section 6.3). Create a function that can vectorize and normalize a document:

$normalizedVector(d) = \frac{V(d)}{|V(d)|}$

* Make sure you don't take `TotalFiles` as a word
* Include all words in the corpus, even if they are not in document (so they will be 0)
* Make sure you sort the words alphabetically
* the value for each term is its TF-IDF weight as calculated by your earlier `tfidf` function.

<!-- grade_id: cell-c9694a0ad4a895c6, points: 0, grade: false, solution: true -->
```python
def normalizedVector(document, index):
    """
    Given a document, vectorize and normalize it
    """
    out = [tfidf(term, document, index) for term in sorted(list(index.keys())[:-1])]
    return euclidean_normalize(out)
    
len(normalizedVector('all_well', NewIndex)), normalizedVector('all_well', NewIndex)[:5]

# Check that the Euclidean length of your vector equals 1
```

Output:
```
(22328, [0.0, 0.0, 0.0, 0.0, 0.0028840469429352866])
```

<!-- grade_id: cell-2e8497964ac87a83, points: 1, grade: true, solution: false -->
```python
assert_equal(len(normalizedVector('all_well', NewIndex)), len(NewIndex)-1)

```

### 9.4 Cosine Similarity

Use the normalized lengths to compute the cosine-similarity between a query and a document. I.e. implement scoring formula 6.12 from the book.

Do this in the form of the function `cosineSimilarity(query, document, index)`

The query vector is a boolean vector with the value for a term being 1 if and only if it occurs at least once in the query.

<!-- grade_id: cell-a87cb124f00cdf2e, points: 0, grade: false, solution: true -->
```python
def cosineSimilarity(query, document, index):
    """
    Given a query and a document, compute and return the cosine-similarity
    between the two.
    """
    qvector = euclidean_normalize([1 if key in query else 0 for key in sorted(list(index.keys())[:-1])])
    dvector = normalizedVector(document, index)
    return np.dot(qvector, dvector)
    
cosineSimilarity(['brutus','love'], 'a_and_c', NewIndex)
```

Output:
```
0.0029940539138864973
```

<!-- grade_id: cell-c960058b3d6d3183, points: 1, grade: true, solution: false -->
```python
assert type(cosineSimilarity(['love','lightly'], 'all_well', NewIndex)) in [float, np.float64, np.float32]
```

### 9.5 Okapi BM25
Now create a search engine which scores according to another famous scoring formula [Okapi BM25](https://en.wikipedia.org/wiki/Okapi_BM25).

<!-- grade_id: cell-236a53b720ae81a9, points: 0, grade: false, solution: true -->
```python

def rankedSearchBM25(query, index, a=1.5, b=0.75):
    """
    Given a query, score the documents according to the BM25 scoring
    scheme and return a list of  (document, score) pairs  in decreasing order
    """
    # this get data on all data for all files
    docinfo = defaultdict(int)
    for word in index.keys():
      if type(index[word]) != int:
        for (doc,num) in  index[word]['Freq_per_Doc'].items():
          docinfo[doc] = docinfo[doc] + num

    doc_avg  = np.average(list(docinfo.values()))

    ans = defaultdict(float)
    for word in query:
      print(word)
      if word not in index.keys():
        continue
      for document in index[word]['Freq_per_Doc']:

          # this the top part of the equation of ranked BM25 search.
          tempValue = tfidf(word,document,index) * (a + 1)
          #this is the down part of the equation
          tempDown = doc_avg / docinfo[document]
          tempDown = index[word]['Freq_per_Doc'][document] + a * (1 - b + b * tempDown)
          ans [document] = ans[document] + tempValue / tempDown


    ans = list(ans.items())
    ans.sort(key = lambda x:x[1], reverse= True)
    return ans

    
rankedSearchBM25(['brutus'], NewIndex)
```

Output:
```
brutus
[('j_caesar', 5.499730851193349),
 ('coriolan', 5.458626859324228),
 ('a_and_c', 4.109740121715352),
 ('hamlet', 2.4946556692411903),
 ('hen_v', 2.3457365256047558),
 ('hen_vi_2', 2.322934341467456),
 ('merchant', 2.1355641732578703),
 ('titus', 2.1157005362512553)]
```

<!-- grade_id: cell-860fbb8c8bd77a2a, points: 1, grade: true, solution: false -->
```python
result = rankedSearchBM25(['brutus'], NewIndex)
assert_equal(type(result),list)
assert_equal(type(result[0]),tuple)
assert_equal(type(result[0][0]),str)
assert isinstance(result[0][1],float)
```

Output:
```
4.109740121715352
5.458626859324228
2.4946556692411903
2.3457365256047558
2.322934341467456
5.499730851193349
2.1355641732578703
2.1157005362512553
```

```python

```